A Deep Learning Pipeline for Mapping in situ Network-level Neurovascular Coupling in Multi-photon Fluorescence Microscopy

Author response:

The following is the authors’ response to the original reviews.

Public Review:

We would like to thank the reviewers for providing constructive feedback on the manuscript. To address their concerns, we have performed additional experiments, analyzed the new data, and revised the manuscript.

(1) The utility of a pipeline depends on the generalization properties.

While the proposed pipeline seems to work for the data the authors acquired, it is unclear if this pipeline will actually generalize to novel data sets possibly recorded by a different microscope (e.g. different brand), or different imagining conditions (e.g. illumination or different imagining artifacts) or even to different brain regions or animal species, etc.

The authors provide a 'black-box' approach that might work well for their particular data sets and image acquisition settings but it is left unclear how this pipeline is actually widely applicable to other conditions as such data is not provided.

In my experience, without well-defined image pre-processing steps and without training on a wide range of image conditions pipelines typically require significant retraining, which in turn requires generating sufficient amounts of training data, partly defying the purpose of the pipeline.

It is unclear from the manuscript, how well this pipeline will perform on novel data possibly recorded by a different lab or with a different microscope.

To address the generalizability of our DL segmentation model, we have performed several validation experiments with deploying our model on out-of-distribution data that 1) had distinct channels 2) were acquired in different species (rat) with a different vascular fluorescent label and a different imaging protocol, and 3) were acquired on a different microscope and with a different vascular label. We first used our model to segment images (507x507um lateral FOV, 170-250 um axial range) from three C57BL/6 mice imaged on the same two-photon fluorescent microscope following the same imaging protocol. The vasculature was labelled by intravenous injection of the Texas Red dextran (70 kDa MW, Thermo Fisher Scientific Inc, Waltham MA), as in the current experiment. In lieu of the EYFP signal from pyramidal neurons that was present in the original data, we added Gaussian noise with a mean and standard deviation identical to the acquired vascular channel in the out-of-distribution dataset. Second, we applied our model to images (507x507um lateral FOV, 300-400 um axial range) from two Fischer rats that were injected with 2000-kDa Alexa680-dextran via a tail vein catheter. These rats were imaged on the same two-photon fluorescence microscope, but with Galvano scanners (instead of resonant scanners). As before, a second channel of Gaussian noise was added to simulate the missing EYFP signal. Finally, we segmented an image of vasculature from an ex-vivo cleared mouse brain (1665x1205x780 um) acquired on a light sheet fluorescence microscope (Miltenyi UltraMicroscope Blaze), with a Lectin-DyLight 649 labelling the vessel walls. The Dice Score, Precision, Recall, Hausdorff 95%, and Mean surface distance were reported for segmentations of 2PFM data sets, following the generation of ground truth images by assisted manual segmentation in ilastik. Examples of the generated segmentation masks are presented in Supplementary figure 9 for visual comparison. We have described the image pre-processing steps/transforms before model inference in the revised Methods section. In general, should the segmentation results on a data set be deemed unsatisfactory, our model can be further fine-tuned on out-of-distribution data. Furthermore, the image analyses downstream from segmentation are applicable irrespective of the method utilized to arrive at a robust vascular segmentation.

(2) Some of the chosen analysis results seem to not fully match the shown data, or the visualization of the data is hard to interpret in the current form.

We have updated the visualizations to make them more accessible and ensure close correspondence between tables and figures.

(3) Additionally, some measures seem not fully adapted to the current situation (e.g. the efficiency measure does not consider possible sources or sinks). Thus, some additional analysis work might be required to account for this.

Thank you for your comment. The efficiency metric was selected as it does not consider sources or sinks. We do agree that accounting for vessel subtypes in the analysis (thus classifying larger vessels as either suppliers/sources or drainers/sinks) would be very useful: notwithstanding, this classification is extremely laborious, as we have noted in our prior work1 . We are therefore leveraging machine learning in a parallel project to afford vessel classification by type. Notwithstanding, the source/sink analysis based on in vivo 2PFM data is confounded by the small FOV.

(4) The authors apply their method to in vivo data. However, there are some weaknesses in the design that make it hard to accept many of the conclusions and even to see that the method could yield much useful data with this type of application. Primarily, the acquisition of a large volume of tissue is very slow. In order to obtain a network of vascular activity, large volumes are imaged with high resolution. However, the volumes are scanned once every 42 seconds following stimulation. Most vascular responses to neuronal activation have come and gone in 42 seconds so each vessel segment is only being sampled at a single time point in the vascular response. So all of the data on diameter changes are impossible to compare since some vessels are sampled during the initial phase of the vascular response, some during the decay, and many probably after it has already returned to baseline. The authors attempt to overcome this by alternating the direction of the scan (from surface to deep and vice versa). But this only provides two sample points along the vascular response curve and so the problem still remains.

We thank the Reviewer for bringing up this important point. Although vessels can show relatively rapid responses to perturbation, vascular responses to photostimulation of ChannelRhodopsin-2 in neighbouring neurons are long-lasting: they do not come and go in 42 seconds. To demonstrate this point, we acquired higher temporal-resolution images of smaller volumes of tissue over 5 minutes preceding and 5 minutes following the 5-s photoactivation with the original photostimulation parameters. The imaging protocol was different in that we utilized a piezoelectric motor, a smaller field of view (512um x (80-128)um x (34-73)um), and only 3x frame averaging, resulting in a temporal resolution of 1.57-3.17 seconds per frame. This acquisition was repeated at different cortical depths in three Thy1-ChR2 mice and the vascular radii were estimated using our presented pipeline. Significantly responding vessels here were selected via an F-test of radius estimates before vs. after stimulation. LOESS fits to the time-dependent radius of significantly responding vessels are shown in Supplementary Figure 5. Vessels shorter than 20 um in length were excluded from the analysis so as to focus on vessel segments where averaging the vascular radius over many vertices was possible. A video of one of the acquisitions is shown along with the timecourses of select vessels’ calibre changes in Author response image 1. The vascular calibre changes following photostimulation persisted for several minutes, consistent with earlier observations by us and others2–5. These small-volume acquisitions demonstrated that dilations were repeatedly longer than the 42 seconds (i.e. our original temporal resolution).

Our temporal sampling was chosen to permit a large field of view acquisition while still being well within the span of the vascular response to look at larger scale vascular coordination that has not previously been studied. The pipeline readily adapts to smaller fields of view at a finer temporal sampling, though such an acquisition precludes the study of the response coordination across hundreds of vessels. While a greater number of baseline frames would help with the baseline variability estimation, maintaining animals under anesthesia during prolonged imaging is exceedingly difficult, precluding us from extending our total acquisition time.

Estimated vascular radius at each timepoint for select vessels from the imaging stack shown in the following video: https://flip.com/s/kB1eTwYzwMJE

(5) A second problem is the use of optogenetic stimulation to activate the tissue. First, it has been shown that blue light itself can increase blood flow (Rungta et al 2017). The authors note the concern about temperature increases but that is not the same issue. The discussion mentions that non-transgenic mice were used to control for this with "data not shown". This is very important data given these earlier reports that have found such effects and so should be included.

We have updated the manuscript to incorporate the data on volumetric scanning in (nontransgenic) C57BL/6 mice undergoing blue light stimulation, with identical parameters as those used in Thy-ChR2 mice (Supplementary Figure 8). As before, responders were identified as vessels that following blue light stimulation showed a radius change greater than 2 standard deviations of their baseline radius standard deviation: their estimated radii changes are shown in Supplementary Figure 8. There was no statistical difference between the radii distributions of any of the photostimulation conditions and pre-photostimulation baseline.

(6) Secondly, there doesn't seem to be any monitoring of neural activity following the photo-stimulation. The authors repeatedly mention "activated" neurons and claim that vessel properties change based on distance from "activated" neurons. But I can't find anything to suggest that they know which neurons were active versus just labeled. Third, the stimulation laser is focused at a single depth plane. Since it is single-photon excitation, there is likely a large volume of activated neurons. But there is no way of knowing the spatial arrangement of neural activity and so again, including this as a factor in the analysis of vascular responses seems unjustified.

Given the high fidelity of Channel-Rhodpsin2 activation with blue light photostimulation found by us and others3, we assume that all labeled neurons within the volume of photostimulation are being activated. Depending on their respective connectivities, their postsynaptic neurons (whether or not they are labeled) may also get activated. We therefore agree with the reviewer that the spatial distribution of neuronal activation is not well defined. The manuscript has been revised to update the terminology from activated to labeled neurons and stress in the Discussion that the motivation for assessing the distance to the closest labeled neuron as one of our metrics is purely to demonstrate the possibility of linking vascular response to activations in their neighbouring neurons and including morphological metrics in the computational pipeline.

(7) The study could also benefit from more clear illustration of the quality of the model's output. It is hard to tell from static images of 3-D volumes how accurate the vessel segmentation is. Perhaps some videos going through the volume with the masks overlaid would provide some clarity. Also, a comparison to commercial vessel segmentation programs would be useful in addition to benchmarking to the ground truth manual data.

We generated a video demonstrating the deep-learning model outputs and have made the video available here: https://flip.com/s/_XBs4yVxisNs. We aimed to develop an open-source method for the research community as the vast majority of groups do not have access to commercial software for vessel segmentation.

(8) Another useful metric for the model's success would be the reproducibility of the vessel responses. Seeing such a large number of vessels showing constrictions raises some flags and so showing that the model pulled out the same response from the same vessels across multiple repetitions would make such data easier to accept.

We have generated a figure demonstrating the repeatability of the vascular responses following photostimulation in a volume and presented them next to the corresponding raw acquisitions for visual inspection (Supplementary figure 6). It is important to note that there is a significant biological variability in vessels’ responses to repeated stimulation, as described previously 3,6: a well-performing model should be able to quantify biological heterogeneity as it of itself may be of interest. Constrictions have been reported in the literature by our group and others 1,2,4,5,7, though their prevalence has not been systematically studied to date. Concerning the reproducibility of our analysis, we have demonstrated model reproducibility (as a metric of its success) on a dataset where vessels visually appeared to dilate consistently following 452 nm light stimulation: these results are now presented in Supplementary Figure 6 of the revised Manuscript. We thus observed that the model repeatedly detected the vessels - that appeared to dilate on visual inspections - as dilating. Examples of vessels constricting repeatedly were also examined and maximal intensity projections of the vessel before and after photostimulation inspected, confirming their repeated constriction (Author response image 2).

It is also worth noting that while the presence of the response (defined as change above 2 standard deviations of the radius across baseline frames) was infrequent (2107 vessels responded at least once, out of a total of 10,552 unique vessels imaged), the direction of the response was highly consistent across trials. Given twice the baseline variability as the threshold for response, of the vessels that responded more than once, 31.7% dilated on some trials while constricting on others; 41.1% dilated on each trial; and 27.2% constricted on each trial. (Note that some trials use 1.1 vs. 4.3 mW/mm2 and some have opposite scanning directions).

Sample capillaries constrictions from maximum intensity projections at repeated time points following optogenetic stimulation.

Baseline (pre-stimulation) image is shown on the left and the post-stimulation image, is on the right, with the estimated radius changes listed to the left.

(9) A number of findings are questionable, at least in part due to these design properties. There are unrealistically large dilations and constrictions indicated. These are likely due to artifacts of the automated platform. Inspection of these results by eye would help understand what is going on.

Some of the dilations were indeed large in magnitude. We present select examples of large dilations and constrictions ranging in magnitude from 2.08 to 10.80 um for visual inspection (Author response image 3) (for reference, average, across vessel and stimuli, the magnitude of radius changes were 0.32 +/- 0.54 um). Diameter changes above 5 um were visually inspected.

Additional views of diameter change in maximum intensity projections ranging in magnitude from 2.

08 um to 10.80 um.

(10) In Figure 6, there doesn't seem to be much correlation between vessels with large baseline level changes and vessels with large stimulus-evoked changes. It would be expected that large arteries would have a lot of variability in both conditions and veins much less. There is also not much within-vessel consistency. For instance, the third row shows what looks like a surface vessel constricting to stimulation but a branch coming off of it dilating - this seems biologically unrealistic.

We now plot photostimulation-elicited vessel-wise radius changes vs. their corresponding baseline radius standard deviations (Author response image 4). The Pearson correlation coefficient between the baseline standard deviation and the radius change was 0.08 (p<1e-5) for 552nm 4.3 mW/mm^2 stimulation, -0.08 (p<1e-5) for 458nm 1.1 mW/mm^2 stimulation, and -0.04 (p<1e-5) for 458nm 4.3 mW/mm^2 stimulation. For non-control (i.e. blue) photostimulation conditions, the change in the radius is thus negatively correlated to the vessel’s baseline radius standard deviation: this small negative correlation indicates that there is little correlation between vessel radius change and the baseline variability in the vessel radius. Classification of vessels by type (arteries vs. veins) is needed before we can comment on differences between these vascular components. The between-vessel (i.e. between parent vessels and their daughter branches separated by branch points) consistency is explicitly evaluated by the assortativity metric, in Figure 9: vessels do somewhat tend to react similarly to their downstream branches: we observed a mean assortativity of 0.4. As for the instance of a surface vessel constricting while a downstream vessel dilates, it is important to remember that the 2PFM FOV restricts us to imaging a very small portion of the cortical microvascular network: one (among many) daughter vessels showing changes in the opposite direction to the parent vessel is not violating the conservation of mass; in addition, mural cells on adjacent branches can respond differently.

Vessel radius change elicited by photostimulation vs.

baseline radius standard deviation across all vessels. The threshold level for response identification is shown as the black line.

(11) As mentioned, the large proportion of constricting capillaries is not something found in the literature. Do these happen at a certain time point following the stimulation? Did the same vessel segments show dilation at times and constriction at other times? In fact, the overall proportion of dilators and constrictors is not given. Are they spatially clustered? The assortativity result implies that there is some clustering, and the theory of blood stealing by active tissue from inactive tissue is cited. However, this theory would imply a region where virtually all vessels are dilating and another region away from the active tissue with constrictions. Was anything that dramatic seen?

The kinetics of the vascular responses are not accessible via the current imaging protocol and acquired data; however, this computational pipeline can readily be adapted to test hypotheses surrounding the temporal evolution of the vascular responses, as shown in Supplementary Figure 2 (with higher temporal-resolution data). Some vessels dilate at some time points and constrict at others as shown in Supplementary Figure 2. As listed in Table 2, 4.4% of all vessels constrict and 7.5% dilate for 452nm stimulation at 4.3 mW/mm^2. There was no obvious spatial clustering of dilators or constrictors: we expect such spatial patterns to be more common with different modes of stimulation and/or in the presence of pathology. The assortativity peaked at 0.4 (quite far from 1 where each vessel’s response exactly matches that of its neighbour).

(12) Why were nearly all vessels > 5um diameter not responding >2SD above baseline? Did they have highly variable baselines or small responses? Usually, bigger vessels respond strongly to local neural activity.

In Author response image 5, we now present the stimulation-induced radius changes vs. baseline radius variability across vessels with a radius greater than 5 um. The Pearson correlation between the radius change and the baseline radius standard deviation across time was low: r=0.05 (p=0.5) for 552nm 4.3 mW/mm^2 stimulation, r=-0.27 (p<1e-5) for 458nm 1.1 mW/mm^2 stimulation, and r=-0.31 (p<1e-5) for 458nm 4.3 mW/mm^2 stimulation. These results demonstrate that the changes following optogenetic stimulation are lower than twice the baseline standard deviation across time for most of these vessels. The pulsatility of arteries results in significant variability in their baseline radius8; in turn, literature to date suggests very limited radius changes in veins. Both of these effects could contribute to the radius response not being detected in many larger vessels.

The change in the vessel radius elicited by photostimulation vs.

baseline vessel radius standard deviation in vessels with a baseline radius greater than 5 um. The threshold level for response identification is shown as the black line.

References

(1) Mester JR, Rozak MW, Dorr A, Goubran M, Sled JG, Stefanovic B. Network response of brain microvasculature to neuronal stimulation. NeuroImage. 2024;287:120512. doi:10.1016/j.neuroimage.2024.120512

(2) Alarcon-Martinez L, Villafranca-Baughman D, Quintero H, et al. Interpericyte tunnelling nanotubes regulate neurovascular coupling. Nature. 2020;kir 2.1(7823):91-95. doi:10.1038/s41586-020-2589-x

(3) Mester JR, Bazzigaluppi P, Weisspapir I, et al. In vivo neurovascular response to focused photoactivation of Channelrhodopsin-2. NeuroImage. 2019;192:135-144. doi:10.1016/j.neuroimage.2019.01.036

(4) O’Herron PJ, Hartmann DA, Xie K, Kara P, Shih AY. 3D optogenetic control of arteriole diameter in vivo. Nelson MT, Calabrese RL, Nelson MT, Devor A, Rungta R, eds. eLife. 2022;11:e72802. doi:10.7554/eLife.72802

(5) Hartmann DA, Berthiaume AA, Grant RI, et al. Brain capillary pericytes exert a substantial but slow influence on blood flow. Nat Neurosci. Published online February 18, 2021:1-13. doi:10.1038/s41593-020-00793-2

(6) Mester JR, Bazzigaluppi P, Dorr A, et al. Attenuation of tonic inhibition prevents chronic neurovascular impairments in a Thy1-ChR2 mouse model of repeated, mild traumatic brain injury. Theranostics. 2021;11(16):7685-7699. doi:10.7150/thno.60190

(7) Hall CN, Reynell C, Gesslein B, et al. Capillary pericytes regulate cerebral blood flow in health and disease. Nature. 2014;508(7494):55-60. doi:10.1038/nature13165

(8) Meng G, Zhong J, Zhang Q, et al. Ultrafast two-photon fluorescence imaging of cerebral blood circulation in the mouse brain in vivo. Proc Natl Acad Sci U S A. 2022;119(23):e2117346119. doi:10.1073/pnas.2117346119

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Line 207: a superfluous '.' before the references.

This has been corrected.

Line 273 ff:

While the metrics are described in mathematical terms which is very useful, the appearing distances (d) and mathematical symbols are not. While mostly intuitively clear, precise definitions of all symbols introduced should be given to avoid ambiguities.

The description has been clarified.

This applies to all formulas appearing in the manuscript and the authors might want to check them carefully.

We have updated them wherever needed.

The mean surface distance seems not to reflect the mean MINIMAL surface distance but just the overall mean surface distance. Or a different definition of the appearing symbols is used, highlighting the need for introducing every mathematical symbol carefully.

The definitions have been updated for clarity, specifying the distinction between Hausdorff 95% distance and mean surface distance.

Line 284:

It is unclear to me why center-line detection was performed in MATLAB and not Python. Using multiple languages/software packages and in addition relying on one that is not freely available/open source makes this tool much less attractive as a real open-source tool for the community. The authors stress in the manuscript abstract that their pipeline is an open and accessible tool, the use of MATLAB defies this logic to some extent in my view.

Centerline detection for large volumetric data is available in Python, see e.g. Scipy packages as well for large data sets via ClearMap or VesselVio.

We tested the centerline detection in Python, scipy (1.9.3) and Matlab. We found that the Matlab implementation performed better due to its inclusion of a branch length parameter for the identification of terminal branches, which greatly reduced the number of false branches; the Python implementation does not include this feature (in any version) and its output had many more such “hair” artifacts. Clearmap skeletonization uses an algorithm by Palagyi & Kuba(1999) to thin segmentation masks, which does not include hair removal. Vesselvio uses a parallelized version of the scipy implementation of Lee et al. (1994) algorithm which does not do hair removal based on a terminal branch length filter; instead, Vesselvio performs a threshold-based hair removal that is frequently overly aggressive (it removes true positive vessel branches), as highlighted by the authors.

Moreover, the authors mention that robust center-line detection was critical. In my view, robust center-line extraction typically requires some additional processing of the binarized data, e.g. using a binary smoothing step. Various binary smoothers are available in the literature and as Python code.

Indeed, binary smoothing was performed: background “holes” located within the vasculature were filled; the masks were dilated (3x) and then eroded to the centreline. Scipy’s binary closing function smoothes the morphology of binary segmentation masks by dilating and then eroding the segmentation masks (as a part of the selected skeletonization algorithm).

Line 303:

'RBC' is not defined (red blood cells?)

This has been updated.

Line 398:

pPhotonsimulation -> Photostimulation

This has been corrected.

Line 400 ff: Efficiency:

I am not sure how useful the measure really is without any information about the 'sources' (i.e. arteries) and sinks (i.e. veins) as blood does not need to be moved between any two arbitrary nodes.

While blood reversals are observed, blood is typically not moved arbitrarily between two arbitrary nodes in capillary networks.

We agree with the reviewer that classifying the vessels by type is important and are currently working on deep learning-based algorithms for the classification of microvasculature into arterioles and venules for future work.

In addition, short paths between two nodes with low resistivity will potentially dominate the sum and the authors excluded vessels 10um and above. This threshold seems arbitrary.

The 10-um diameter threshold was not applied in the computation of the network metrics. The 10-um thresholding was restricted to “capillary” identification in Figure 8: the 10-um cutoff for referring to a vessel as a capillary has long been applied in the literature [1], [2], [3], [4], [5], [6], [7], [8], [9], [10], [11].

Figure 3:

It's unclear what the units are for the Mean Surface and Harsdorf Distances (pixel or um?).

The units have now been specified (um).

Figure 4:

The binarized data, and particularly the crops are difficult to interpret in black and white. It would be much more useful to present the segmentation results in a way that is interpretable (e.g. improving the rendering of the 3d information, particularly in the crops by using shadows or color codes for depth, etc).

We have updated these visualizations and shaded them based on cortical depth.

Panel C indicates that the illastik is performing badly due to changes in imagining conditions (much higher background level). As pointed out before, in my view, a reasonable pipeline should start by removing and standardizing background levels as well as dynamic ranges and possibly other artifacts before performing a more detailed analysis. This would also make the pipeline more robust against data from other microscopes etc as only a few preprocessing parameters might need to be adjusted.

I wonder whether after such a pre-processing step, UNET / UNETR would still perform in a way that was superior to ilastik, as ground truth data was generated with the aid of illastiks initially.

The Ilastik model is based on semi-automatically generated foreground labels in small batches. We had to break it up into small groups during manual labelling as larger groups were not able to run due to the computational limits of Ilastik. Ilastik is typically trained in an iterative fashion on a few patches at a time because it takes 2-3 hours per patch to train and the resulting model does not generalize on the remaining patches or out-of-distribution data - even with image pre-processing steps. On the reviewer's comment, we did try inputting normalized images into Ilastik, but this did not improve its results. UNET and UNETR inputs have been normalized for signal intensities.

Typical pre-processing/standard computer vision techniques with parameter tuning do not generalize on out-of-distribution data with different image characteristics, motivating the shift to DL-based approaches.

Figure 5:

This is a validation figure that might be better shown in an appendix or as a supplement.

Since this is a methodological paper, we think it is important to highlight the validation of the proposed method.

Line 476:

It's surprising that the number of vessel segments almost doubles when taking the union. Is the number of RBC plugs expected to be so high?

The etiology of discontinuities includes, but is not limited to, RBC plugs; we expect discontinuities to arise also from a very short pixel dwell time (0.067us) of the resonant scanning and have indeed observed apparent vessel discontinuities on resonant scanning that are not present with Galvano scanning using a pixel dwell time of 2us.

Section 4.4 / 4.5 :

The analysis in these sections provides mostly tables with numbers that are more difficult to read and hides possible interesting structures in the distribution of the various measures/quantities. For example, why is 5um a good choice to discriminate between small and large vessels, why not resolve this data more precisely via scatter plots?

Some distributions are shown in the appendix and could be moved to the main analysis.

Generally, visualizing the data and providing more detailed insights into the results would make this manuscript more interesting for the general reader.

The radius of vessel segments drops off after 5.0 um, as shown in Supplementary Figure 4A. The 10-um diameter thresholding is based on prior literature [1], [12], [13], [14], [15], [16], [17], [18], [19] and is used to segregate different vessel types in a conservative manner. The smallest capillaries are expected to have pericytes on their vessel walls whereas arteries are expected to have smooth muscle cells on their vessel walls. These differences in mural cells also may lead to differences in respective vessels’ reactivity.

The data summarized in Tables 1 and 2 are shown as scatter plots in Figures 8, Supplementary Fig 4 and Supplementary Fig 5.

Line 556:

The authors deem a certain change in radius as the relevant measure for responding vessels. They deem a vessel responding if it dilates by twice the std deviation in the radius.

Based on this measure they find that large vessels rarely respond.

However, I think this analysis might obscure some interesting effects:

(1) The standard deviation of the radius depends on the correct estimation of the center point. Given the limited spatial resolution the center point (voxel) obtained from the binarization and skeletonization might not lie in the actual center of the vessel. This effect will be stronger for larger vessels. Center point coordinates should thus be corrected to minimize the std in radius.

(2) Larger vessels will not necessarily have a perfectly circular shape, and thus the std measure is not necessarily a good measure of 'uncertainty' of estimating the actual radius.

(3) The above reasons possibly contribute to the fact that from Figure 6 it seems vessels with larger radii have higher std in general (as indicated above some more detailed visualization of the data instead of plain tables could reveal such effects better, e.g. scatter radius vs std). This higher std is making it harder to detect changes in larger vessels. However, with respect to the blood flow, the critical factor is the cross-section of the vessel that scales with the radius squared. Thus, a fixed change in radius for a vessel (say 1um) will induce a larger increase in the flow rate in larger vessels as the change in cross-section is also proportional to the radius of the vessel.

Thus, larger vessels to be deemed responders should probably have lower thresholds, thresholds should be taken on the cross-section change, or at least thresholds should not be higher for larger vessels as it is the case now using the higher std.

(1) The radius estimate does not depend on the precise placement of the center point as the radius is not being estimated by the distance from the center point to the boundary of the vessel. Instead, our strategy is to estimate the cross-sectional area (A) of the vessel by the Riemann sum of the sectors with the apex at the center point; the radius is then quoted as sqrt(A/pi) (Supplementary figure 3B). Thus, estimated vessel radius estimates in each cross-sectional plane are then averaged across the cross-sectional planes placed every ~1um along the vessel length. The uncertainty in the cross-sectional plane’s vessel radius, the uncertainty in the vessel radius (upon averaging the cross-sectional planes), and the uncertainty in the radius estimate across repeated measures of a state (i.e. across different samples of the baseline vs, post-photostimulation states) are all reported, and the last one used to define responding vessels.

To demonstrate the insensitivity to the precise placement of the vessel’s centrepoint, we have jittered the centerline in the perpendicular plane to the vessel tangent plane at each point along the vessel and then estimated the mean radius in 71 cross-sectional planes of larger vessels (mean radius > 5 um). The percent difference in the estimated radius at our selected vessel centrepoints vs. the jittered centrepoints is plotted above. The percent difference in the mean radius estimated was 0.64±3.44% with 2.45±0.30 um centerpoint jittering. (In contrast, photostimulation was estimated to elicit an average 25.4±18.1% change in the magnitude of the radius of larger vessels, i.e. those with a baseline radius >5um.)

(2) Indeed, the cross-sectional areas of either large or small vessels are not circles. Consequently, we are placing the vessel boundary, following other published work[20], at the minimum of the signal intensity gradients computed along thirty-six spokes emanating from the centrepoint (cf Figure 2H,K). The cross-sectional area of the vessel in the said cross-sectional plane is then estimated by summing the areas of the sectors flanked by neighbouring spokes. We do not make an assumption about the cross-sectional area being circular. We report radii of circles with the equivalent area as that of the cross-sectional areas merely for ease of communication (as most of the literature to date reports vessel radii, rather than vessel cross-sectional areas.)

To demonstrate the robustness of this approach, we show the sensitivity of vessel-wise radius estimate on the number of spokes used to estimate the radius in Supplementary Figure 3a. The radius estimate converges after 20 spokes have been used for estimation. Our pipeline utilizes 36 spokes and then excludes minima that lie over 2 STD away from the mean radius estimate across those 36 spokes. With 36 spokes, the vesselwise mean radius estimation was within 0.24±0.62% of the mean of radius estimates using 40-60 spokes.

(3) Across-baseline sample uncertainty in vessel radius is not dependent on baseline vessel caliber (i.e. this uncertainty is not larger in larger vessels).

Supplementary Figure 5 shows vessel radius changes for large vessels without a threshold defining responding or non-responding vessels. To explore the dependence of the outcomes on the threshold used to identify the responding vessels, we have explored an alternative strategy, whereby responding small vessels are identified as those vessels that show a post-photostimulation (vs. baseline) radius change of more than 10%. These data are now plotted in Supplementary Figure 10, for capillaries which is in agreement with Figure 8. These points are now also discussed in the Discussion section of the revised manuscript:

“Additionally, alternative definitions of responding vessels may be useful depending on the end goal of a study (e.g., this could mean selecting a threshold for the radius change based on a percentage change from the baseline level).”

Section 4.5.1

Why is the distance to the next neuron a good measure here? If two or more neurons are just a bit further away there will be twice or multiple times the 'load' while the measure would only indicate the distance to the shortest neuron. I wonder how the results change if those 'ensemble' effects are taken into account.

In this direction, looking for network-level effects with respect to the full spatial organization of the neurons would be very interesting to look at.

We agree with the review that this question is interesting; however, it is not addressable using present data: activated neuronal firing will have effects on their postsynaptic neighbors, yet we have no means of measuring the spread of activation using the current experimental model.

Figure 8

The scatter plots shown are only partly described (e.g. what's the line with error bars in C, why does it only appear for the high-intensity stimulation?).

Quadratic polynomial fit is shown only in C as the significant response was observed only for this condition, i.e. for the higher intensity blue photostimulation.

From the scatter plots as shown it is not clear to me why dilations happen on average further away. This might be a density effect not well visible in this representation. The data does not seem to show a clear relationship between neuron distance and Delta R.

Particularly in the right panel (high stimulation) there seems to be a similar number of close by neurons responding in both directions, but possibly a few more contracting at larger distances?

So, the overall effect does not seem as 'simple' as suggested in the title of section 4.5.1 in my view, but rather more cells start to contract at larger distances while there seems to be a more intricate balance nearby.

A more thorough analysis and visualization of the densities etc. might be needed to clarify this point.

The language has been revised to:

458-nm photostimulation resulted in a mix of constrictions and dilations with 44.1% of significantly responding vessels within 10 um of a labelled pyramidal neuron constricting and 55.1% dilating, while 53.3% of vessels further than 30 um constricted and 46.7% dilated. The cutoff distances from the closest labelled neuron were based on estimates of cerebral metabolic rate of oxygen consumption that showed a steep gradient in oxygen consumption with distance from arteries, CMRO2 being halved by 30 μm away

We added a probability density plot for significant constrictors and dilators to Figure 8 and Supplementary Figure 5.

Figure 8 Panel D / Section 4.5.2

This is a very interesting result in my view found in this study.

I am unclear how to interpret the effect. The authors state that dilators tend to be closer to the surface. Looking at the scatter plot (without real density information except the alpha value) it seems again the number of responders in both directions is about the same, but in deeper regions the contraction is just larger? This would be different, than how the authors interpret the data. It is unclear from the provided analysis/plots what is actually the case.

We added a probability density function plot of the constrictors and dilators, which shows a greater incidence of constrictions (vs. dilations). The text of the paper was then clarified to include the proportion of significant constrictors/ dilators closer than 10 um vs. further than 30 um away from the closest labeled neuron.

For the analyses above involving $Delta R$ I recommend also look how those results change when looking at changes in cross section instead, i.e. taking into account the actual vessel radius as well as discussed above.

It would be interesting to speculate here or in the discussion on a reason why vessels in deeper regions might need to contract more?

Unaddressed is the question if e.g. contraction in a vessel for small stimulation is predictive of contractions for larger stimulation or any other relationships?

Thank you for your comment. Given its hierarchical organization and high within-vessel response heterogeneity, we believe that the vasculature is best analyzed as a network. Our radius estimates come from averaged cross-sectional estimates allowing us to examine heterogeneity within individual vessel segments.

The discussion has been updated to include reasons as to why deeper vessels may contract more:

“As the blue light stimulation power increased, the mean depth of both constricting and dilating vessels increased, likely resulting from higher intensity light reaching ChR2-expressing neurons deeper in the tissue and exciting superficial neurons (and thus their postsynaptic neurons) to a greater level [21], [22]. The blue light would be expected to excite a lower number of neurons farther from the cortical surface at lower powers.”

Also, how consistent are contractions/dilations observed at a particular vessel etc.

To look at the consistency of a particular vessel's response to the 1.1 or 4.3 mW/mm^2 blue light photostimulation, we categorized all significant responses as constrictions or dilations, defining a responding vessel as that showing a change that is either > 2 x baseline vessel radius variability or >10% of the vessel’s mean baseline radius.

Given twice the baseline variability as the threshold for response, of the vessels that responded more than once, 31.7% dilated on some trials while constricting on others; 41.1% dilated on each trial; and 27.2% constricted on each trial. (Note that some trials use 1.1 vs. 4.3 mW/mm2 and some have opposite scanning directions).

Section 4.5.3

The results in assortativity are interesting. It would be interesting to look at how the increase in assortativity is mediated. For, example, is this in localized changes in some parts of the graph as visible in A or are there other trends? Do certain sub-graphs that systematically change their radius have certain properties (e.g. do activated neurons cluster there) or are these effects related to some hotspots that also show a coordinated change in control conditions (the assortativity seems not zero there)?

I already discussed if the efficiency measure is necessarily the best measure to use here without taking into account 'sources' and 'sinks'.

We plan to address this in future work once we have successfully trained models for the classification of vessels into arteries, veins, and capillaries. Capillaries will be classified based on their branch order from parent arteries to specify where in the network changes are occurring.

Figure 9

It's unclear to me why the Ohm symbol needs to be bold?

It is not bolded (just the font’s appearance).

Line 707:

"458-nm photostimulation caused capillaries to dilate when pyramidal neurons were close, and constrict when they were further away."

In my view, this interpretation is too simple, given the discussion above. A more detailed analysis could clarify this point.

The discussion on this point has been revised to:

458-nm photostimulation resulted in a mix of constrictions and dilations, with 44.1% of significantly responding vessels within 10 μm of a labelled pyramidal neuron constricting, and 55.1% dilating; while 53.3% of vessels further than 30 μm constricted and 46.7% dilated. The cutoff distances from the closest labelled neuron were based on estimates of cerebral metabolic rate of oxygen consumption that showed a steep gradient in oxygen consumption with distance from arteries, CMRO2 being halved by 30 μm away [23].

Line 740:

"The network efficiency here can be thought of as paralleling mean transit time, i.e., the time it takes blood to traverse the capillary network from the arteries to the veins".

The network efficiency as defined by the authors seems not to rely on artery/vein information and thus this interpretation is not fully correct in my view.

The authors might want to reconsider this measure for one that accounts for sources and sinks, if they like to interpret their results as in this line.

Yes, the efficiency described does not account for sources and sinks. It estimates the resistivity of capillaries, as a proxy for the ease of moving through the observed capillary nexus. Looking at the efficiency metric from graph theory does not require knowledge of the direction of blood flow, and can comment on the resistivity changes across capillary networks.

For future work, we are investigating methods of classifying vessels as arteries, capillaries, or veins. This type of analysis will provide more detailed information on paths between arteries and veins; it will not provide insight into large-scale network-wide modifications, as those require larger fields of view.

Line 754 Pipeline Limitations and Adaptability

I think the additional 'problem' of generating new training data for novel data sets or data from other microscopes etc should be addressed or the pipeline tested on such data sets.

Generating training data is typically the biggest time investment when adapting pipelines.

The generalization properties of the current pipeline are not discussed (e.g. performance on a different microscope / different brain area / different species etc.).

The public response to reviews has been updated with out-of-distribution data from other imaging protocols, microscopes, and species showing generalizability. These results have also been added to the paper as Supplementary Table 4, and Figure 6. The performance of our pipeline on these out-of-distribution data is now discussed in the updated Discussion section.

Line 810

Code availability should be coupled with the publication of this paper as it seems the main contribution. I don't see how the code can be made available after publication only. It should be directly available once the manuscript is published and it could help to make it available to the reviewers before that. It can be updated later of course.

The code is being made available.

Reviewer #2 (Recommendations For The Authors):

This analytical pipeline could be quite useful but it needs to be better demonstrated. If faster volumetric imaging is not possible, perhaps using it over a small volume would still demonstrate its utility at a smaller but more believable scale.

The higher temporal resolution scans (over smaller tissue volumes) have now been performed and the results of applying our pipeline to these data are summarized in Supplementary Figure 2.

Using sensory stimuli for neuronal activation might be a better idea than optogenetic stimulation. It isn't necessary but it would avoid the blue light issue.

The pipeline is readily applicable for analysis of vasoreactivity following different perturbers; however, the robustness of vessels’ response is higher with blue light photostimulation of ChR2 than with sensory stimuli [24]. Notwithstanding, an example of the vascular response to electrical stimulation of the contralateral forepaw is now included in Supplementary Figure 2.

This tool could be quite useful even without neural activity mapping. It obviously makes it even more powerful, but again, the utility could be demonstrated with just vascular data or even anatomical neuronal data without function.

We agree with both points, and have emphasized them in the revised discussion section.

Line 559 says the average capillary diameter change was 1.04 um. The next sentence and the table below all have different values so this is unclear.

The wording was updated to make this clearer.

Line 584 - should 458 be 552?

458 is correct.

Figure 1 - the schematic doesn't seem right - the 650 LPF with the notches is positioned to pass short light and reflect long wavelengths and the notch bands.

The figure has been updated to reflect this. The original layout was done for compactness.

References

(1) D. A. Hartmann, V. Coelho-Santos, and A. Y. Shih, “Pericyte Control of Blood Flow Across Microvascular Zones in the Central Nervous System,” Annu. Rev. Physiol., vol. 84, no. Volume 84, 2022, pp. 331–354, Feb. 2022, doi: 10.1146/annurev-physiol-061121-040127.

(2) J. Batista, “An adaptive gradient-based boundary detector for MRI images of the brain,” in 7th International Conference on Image Processing and its Applications, Manchester, UK: IEE, 1999, pp. 440–444. doi: 10.1049/cp:19990360.

(3) Y. Le, X. Xu, L. Zha, W. Zhao, and Y. Zhu, “Tumor boundary detection in ultrasound imagery using multi-scale generalized gradient vector flow,” J. Med. Ultrason., vol. 42, no. 1, pp. 25–38, Jan. 2015, doi: 10.1007/s10396-014-0559-3.

(4) X. Ren, “Multi-scale Improves Boundary Detection in Natural Images,” in Computer Vision – ECCV 2008, D. Forsyth, P. Torr, and A. Zisserman, Eds., Berlin, Heidelberg: Springer, 2008, pp. 533–545. doi: 10.1007/978-3-540-88690-7_40.

(5) C. Grigorescu, N. Petkov, and M. A. Westenberg, “Contour and boundary detection improved by surround suppression of texture edges,” Image Vis. Comput., vol. 22, no. 8, pp. 609–622, Aug. 2004, doi: 10.1016/j.imavis.2003.12.004.

(6) J. Tang and S. T. Acton, “Vessel Boundary Tracking for Intravital Microscopy Via Multiscale Gradient Vector Flow Snakes,” IEEE Trans. Biomed. Eng., vol. 51, no. 2, pp. 316–324, Feb. 2004, doi: 10.1109/TBME.2003.820374.

(7) J. Merkow, A. Marsden, D. Kriegman, and Z. Tu, “Dense Volume-to-Volume Vascular Boundary Detection,” in Medical Image Computing and Computer-Assisted Intervention - MICCAI 2016, S. Ourselin, L. Joskowicz, M. R. Sabuncu, G. Unal, and W. Wells, Eds., Cham: Springer International Publishing, 2016, pp. 371–379. doi: 10.1007/978-3-319-46726-9_43.

(8) F. Orujov, R. Maskeliūnas, R. Damaševičius, and W. Wei, “Fuzzy based image edge detection algorithm for blood vessel detection in retinal images,” Appl. Soft Comput., vol. 94, p. 106452, Sep. 2020, doi: 10.1016/j.asoc.2020.106452.

(9) M. E. Martinez-Perez, A. D. Hughes, S. A. Thom, A. A. Bharath, and K. H. Parker, “Segmentation of blood vessels from red-free and fluorescein retinal images,” Med. Image Anal., vol. 11, no. 1, pp. 47–61, Feb. 2007, doi: 10.1016/j.media.2006.11.004.

(10) A. M. Mendonca and A. Campilho, “Segmentation of retinal blood vessels by combining the detection of centerlines and morphological reconstruction,” IEEE Trans. Med. Imaging, vol. 25, no. 9, pp. 1200–1213, Sep. 2006, doi: 10.1109/TMI.2006.879955.

(11) A. F. Frangi, W. J. Niessen, K. L. Vincken, and M. A. Viergever, “Multiscale vessel enhancement filtering,” in Medical Image Computing and Computer-Assisted Intervention — MICCAI’98, W. M. Wells, A. Colchester, and S. Delp, Eds., Berlin, Heidelberg: Springer, 1998, pp. 130–137. doi: 10.1007/BFb0056195.

(12) K. Bisht et al., “Capillary-associated microglia regulate vascular structure and function through PANX1-P2RY12 coupling in mice,” Nat. Commun., vol. 12, no. 1, p. 5289, Sep. 2021, doi: 10.1038/s41467-021-25590-8.

(13) Y. Wu et al., “Quantitative relationship between cerebrovascular network and neuronal cell types in mice,” Cell Rep., vol. 39, no. 12, p. 110978, Jun. 2022, doi: 10.1016/j.celrep.2022.110978.

(14) T. Kirabali et al., “The amyloid-β degradation intermediate Aβ34 is pericyte-associated and reduced in brain capillaries of patients with Alzheimer’s disease,” Acta Neuropathol. Commun., vol. 7, no. 1, p. 194, Dec. 2019, doi: 10.1186/s40478-019-0846-8.

(15) X. Ren et al., “Linking cortical astrocytic neogenin deficiency to the development of Moyamoya disease–like vasculopathy,” Neurobiol. Dis., vol. 154, p. 105339, Jul. 2021, doi: 10.1016/j.nbd.2021.105339.

(16) J. Steinman, M. M. Koletar, B. Stefanovic, and J. G. Sled, “3D morphological analysis of the mouse cerebral vasculature: Comparison of in vivo and ex vivo methods,” PLOS ONE, vol. 12, no. 10, p. e0186676, Oct. 2017, doi: 10.1371/journal.pone.0186676.

(17) A.-A. Berthiaume et al., “Dynamic Remodeling of Pericytes In Vivo Maintains Capillary Coverage in the Adult Mouse Brain,” Cell Rep., vol. 22, no. 1, pp. 8–16, Jan. 2018, doi: 10.1016/j.celrep.2017.12.016.

(18) S. Katz, R. Gattegno, L. Peko, R. Zarik, Y. Hagani, and T. Ilovitsh, “Diameter-dependent assessment of microvascular leakage following ultrasound-mediated blood-brain barrier opening,” iScience, vol. 26, no. 6, p. 106965, Jun. 2023, doi: 10.1016/j.isci.2023.106965.

(19) J. Drouin-Ouellet et al., “Cerebrovascular and blood-brain barrier impairments in Huntington’s disease: Potential implications for its pathophysiology,” Ann. Neurol., vol. 78, no. 2, pp. 160–177, Aug. 2015, doi: 10.1002/ana.24406.

(20) K. P. McDowell, A.-A. Berthiaume, T. Tieu, D. A. Hartmann, and A. Y. Shih, “VasoMetrics: unbiased spatiotemporal analysis of microvascular diameter in multi-photon imaging applications,” Quant. Imaging Med. Surg., vol. 11, no. 3, pp. 969–982, Mar. 2021, doi: 10.21037/qims-20-920.

(21) E. L. Johnson et al., “Characterization of light penetration through brain tissue, for optogenetic stimulation.” bioRxiv, p. 2021.04.08.438932, Apr. 08, 2021. doi: 10.1101/2021.04.08.438932.

(22) S. I. Al-Juboori, A. Dondzillo, E. A. Stubblefield, G. Felsen, T. C. Lei, and A. Klug, “Light scattering properties vary across different regions of the adult mouse brain,” PloS One, vol. 8, no. 7, p. e67626, 2013, doi: 10.1371/journal.pone.0067626.

(23) P. Mächler et al., “Baseline oxygen consumption decreases with cortical depth,” PLOS Biol., vol. 20, no. 10, p. e3001440, Oct. 2022, doi: 10.1371/journal.pbio.3001440.

(24) J. R. Mester et al., “In vivo neurovascular response to focused photoactivation of Channelrhodopsin-2,” NeuroImage, vol. 192, pp. 135–144, May 2019, doi: 10.1016/j.neuroimage.2019.01.036.

Changes from the earlier version to the revised version

Abstract

Abstract

Functional hyperaemia is a well-established hallmark of healthy brain function, whereby local brain blood flow adjusts in response to a change in the activity of the surrounding neurons. Although functional hyperemia has been extensively studied at the level of both tissue and individual vessels, vascular network-level coordination remains largely unknown. To bridge this gap, we developed a deep learning-based computational pipeline that uses two-photon fluorescence microscopy images of cerebral microcirculation to enable automated reconstruction and quantification of the geometric changes across the microvascular network, comprising hundreds of interconnected blood vessels, pre and post-activation of the neighbouring neurons. The pipeline’s utility was demonstrated in the Thy1-ChR2 optogenetic mouse model, where we observed network-wide vessel radius changes to depend on the photostimulation intensity, with both dilations and constrictions occurring across the cortical depth, at an average of 16.1±14.3 μm (mean±stddev) away from the most proximal neuron for dilations; and at 21.9±14.6 μm away for constrictions. We observed a significant heterogeneity of the vascular radius changes within vessels, with radius adjustment varying by an average of 24 ± 28% of the resting diameter, likely reflecting the heterogeneity of the distribution of contractile cells on the vessel walls. A graph theory-based network analysis revealed that the assortativity of adjacent blood vessel responses rose by 152 ± 65% at 4.3 mW/mm2 of blue photostimulation vs. the control, with a 4% median increase in the efficiency of the capillary networks during this level of blue photostimulation in relation to the baseline. Interrogating individual vessels is thus not sufficient to predict how the blood flow is modulated in the network. Our computational pipeline, to be made openly available, enables tracking of the microvascular network geometry over time, relating caliber adjustments to vessel wall-associated cells’ state, and mapping network-level flow distribution impairments in experimental models of disease.

Introduction

Earlier version: To support healthy brain functioning, the cerebrovascular network undergoes continual adjustments in vessel calibers [1]–[3]. Neurovascular coupling refers to the change in blood flow following changes in the level of neuronal activity: under physiological conditions, a generous buffer of nutrients is granted to activated parenchyma via the capillary network [1], [4]. This buffer is maintained through finely tuned regulation of flow through changes in the vessel caliber, mediated via contractile cells in the vessel walls. In the absence of such tuning, pockets of tissue could experience inadequate access to metabolites [5]. Alterations in smooth muscle cells, pericytes, and astrocytes may lead to compromises in vessels’ dilatory capacity and thus deficits in neurovascular coupling [6]–[9]. In various brain pathologies, including Alzheimer’s disease, stroke, and trauma, regional blood flow regulation gets impaired through vessel loss and/or dysfunction of the vessels’ dilatory capacity, resulting in regions of ischemia/hypoxia [10], [11]. Previous studies have examined either individual vessels or the tissue level responses, with little attention having been paid to the vascular network, though network dysfunction frequently is associated with accelerated disease progression and long-term symptomatology [3], [12]–[20].

Revised version: To support healthy brain functioning, the cerebrovascular network undergoes continual adjustments in vessel calibers [1], [2], [3]. Neurovascular coupling refers to the change in blood flow following changes in the level of neuronal activity: under physiological conditions, a generous buffer of nutrients is granted to activated parenchyma via the capillary network [1], [4]. This buffer is maintained through finely tuned regulation of flow through changes in the vessel caliber, mediated via contractile cells in the vessel walls. In the absence of such tuning, pockets of tissue could experience inadequate access to metabolites [5]. Alterations in smooth muscle cells, pericytes, and astrocytes may lead to compromises in vessels’ dilatory capacity and thus deficits in neurovascular coupling [6], [7], [8], [9]. In various brain pathologies, including Alzheimer’s disease, stroke, and trauma, regional blood flow regulation gets impaired through vessel loss and/or dysfunction of the vessels’ dilatory capacity, resulting in regions of ischemia/hypoxia [10], [11]. Previous studies have examined either individual vessels or the tissue level responses, with little attention having been paid to the vascular network, though network dysfunction frequently is associated with accelerated disease progression and long-term symptomatology [3], [12], [13], [14], [15], [16], [17], [18], [19], [20].

Earlier version: While there is copious data on the functioning of individual vessels, interrogation of the microvascular network remains a challenge, in terms of both data acquisition and analysis [6], [7], [21]–[23]. To date, studies on the brain vasculature have been done by sparsely imaging individual blood vessels at the cellular scale [6], [7], [21], [24]–[32], thereby severely undersampling the microvascular network; or by evaluating the averaged flow over many vessels at the mesoscopic scale, thus failing to discern the flow through individual vessels. A critical gap in the field is the characterization of flow across hundreds of individual vessels, while imaging the network structure that links them together to determine how the vascular response is coordinated across the network. This gap is particularly significant as studies investigating blood flow across several vessels at a time (imaged by varying the line acquisition pattern [6], [24]) have shown highly heterogeneous responses among capillaries. Neuronal function impairments arise wherever local metabolite supply becomes inadequate, notwithstanding the physiological level of flow across the network as a whole, making mapping of vessel changes across the network of particular importance.

Revised version: While there is copious data on the functioning of individual vessels, interrogation of the microvascular network remains a challenge, in terms of both data acquisition and analysis [6], [7], [21], [22], [23]. To date, studies on the brain vasculature have been done by sparsely imaging individual blood vessels at the cellular scale [6], [7], [21], [24], [25], [26], [27], [28], [29], [30], [31], [32], thereby severely undersampling the microvascular network; or by evaluating the averaged flow over many vessels at the mesoscopic scale, thus failing to discern the flow through individual vessels. A critical gap in the field is the characterization of flow across hundreds of individual vessels, while imaging the network structure that links them together to determine how the vascular response is coordinated across the network. This gap is particularly significant as studies investigating blood flow across several vessels at a time (imaged by varying the line acquisition pattern [6], [24]) have shown highly heterogeneous responses among capillaries. Neuronal function impairments arise wherever local metabolite supply becomes inadequate, notwithstanding the physiological level of flow across the network as a whole, making mapping of vessel changes across the network of particular importance.

To address the limitations of previous work, we developed a novel deep learning (DL)-based pipeline for mapping changes to the geometry of the brain vascular network following neuronal activation, from a time series of volumetric two-photon fluorescence microscopy (2PFM) data. Neuronal activation was elicited by photostimulation of pyramidal neurons expressing Channelrhodopsin-2 (ChR2) [33] in the Thy1-ChR2-YFP mouse model [34]. Our DL pipeline enabled automatic and accurate segmentation, registration and network analysis of large 2PFM datasets across time. We applied our pipeline in a dataset of 17 Thy1-ChR2-YFP mice to map photostimulation-induced changes across the microvascular network - at the level of individual vessels and at the level of vertices spaced every micrometre along the vessels – in relation to the distance to the closest pyramidal neurons expressing the optogenetic actuator and across the cortical depth. Our findings demonstrate the utility of our pipeline for studying in situ microvascular morphology and function to address various neuroscientific hypotheses.

Methods

Cohort

Animals

All experimental procedures in this study followed the ARRIVE 2.0 guidelines [35]. They were approved by the Animal Care Committee of the Sunnybrook Research Institute, which adheres to the policies and guidelines of the Canadian Council on Animal Care and meets all the requirements of the Provincial Statute of Ontario, Animals for Research Act, and the Canadian Federal Health of Animals Act. We used 15 male and 13 female Thy1-ChR2-YFP mice (#007612, line 18, Jackson Laboratory) at 6-12 months of age (283.9 ± 67.0 days), weighing 28.8±7.1g on the day of imaging. The mice were bred in-house and housed under a 12-hr light/dark cycle [34]. In this mouse line, Channelrhodopsin-2 is expressed preponderantly by pyramidal neurons with soma in cortical layer 5 and dendrites projecting to the cortical surface, enabling depolarization of their cellular membranes and action potential generation upon blue light illumination [30], [34]. An attrition schematic is provided in Supplementary Figure 1.

Surgical Preparation

On the day of imaging, mice were induced with 5% isoflurane, transferred to a rectal probe feedback-regulated heating pad (CWE Inc, Ardmore PA) to maintain a temperature of 37 °C, and maintained under 2% isoflurane. A subcutaneous injection of 1 mL of lactated Ringer’s solution was administered at the start of surgery. Throughout the surgical preparation and imaging, we monitored breath rate, heart rate, arterial blood oxygen saturation, pulse distention, and breath distention via a pulse oximeter, with a probe mounted on the thigh (MouseOx, STARR Life Sciences) (Supplementary Table 1). For fine control over respiration, the mice were tracheostomized with an endotracheal tube (20 Ga catheter) and ventilated with a gas mixture of 20-30% O2 and medical air using a small animal ventilator (SAR 830/P, CWE Inc, Ardmore PA) set to 115-130 breaths per minute at an inspiratory/expiratory ratio of 1:3-4. Their heads were immobilized via ear bars during the placement of the cranial window and imaging. The tail vein was cannulated with a 26 Ga catheter to enable fluorophore and anesthetic delivery. A cranial window was implanted over the forelimb region of the primary somatosensory cortex (AP 0.25mm, ML 2.0 mm); the dura was excised, and 1% agarose was applied between the brain and the glass coverslip. A well was built using dental cement to allow water immersion of the objective. Texas Red dextran (70 kDa MW, Thermo Fisher Scientific Inc, Waltham MA) was diluted in PBS and injected through the tail vein catheter at a concentration of 33 mg/kg for a total injection volume between 0.1 and 0.2 mL. A line containing 0.01 g/ml alpha chloralose was connected to the tail vein catheter. The animal was then positioned underneath the microscope objective, and the structural 2PFM scan acquired, as detailed below. Seventeen mice (nine male and eight female) made it through the surgical preparation (cf. Supplementary Figure 1). Following the structural acquisition, the isoflurane was discontinued, and the continuous infusion of alpha chloralose commenced at 40 mg/kg/hr.

Imaging

Two-photon Fluorescence Imaging and Optogenetic Stimulation

Mice were imaged on an FVMPE-RS microscope (Olympus, Japan) using a 25x/1.05NA objective. An Insight tunable Ti:Saphire near-infrared laser (SpectraPhysics, USA) was used for 900 nm excitation of Texas Red-labelled vasculature and YFP-labeled ChR2-expressing pyramidal neurons. Two visible light stimulation lasers were used to excite ChR2 at either 458 nm or 552 nm (control). The optical setup is shown in Figure 1. The imaging field-of-view was selected to include at least one penetrating artery and vein yet avoid major pial vessels and thus signal loss in the underlying cortex. Structural scans were acquired under 2% isoflurane using a Galvano scanner, with 2x averaging, a z-step of 0.99 μm, and a nominal lateral resolution of 0.99 μm. Following structural scanning, mice were allowed to rest for 5 minutes to let blood flow equilibrate from brain exposure to room light. The five-minute resting period was observed due to prior reports of red blood cell velocities following optogenetic ChR2 stimulation in the same mouse strain remaining altered for a minute following photostimulation and due to potential pericyte-mediated responses via cytoskeleton reorganization occurring over a minute [6], [30].

Photostimulation setup:

The excitation and stimulation light pass through a FV30-NDM690 dichroic mirror with two notch filters, at 458 nm and 552 nm, to excite TexasRed, EYFP and ChR2 within the mouse. The emitted light passes through the objective, is reflected off the FV30-NDM690 dichroic mirror, and passes through a 650 nm barrier filter before reaching a 570 nm long pass filter (LPF) separating emitted light from EYFP and TexasRed, which respectively pass through 495-540 nm and 575-630 nm barrier filters to be collected via GaAsP detectors.

Earlier version: Functional scans were acquired over baseline and post-stimulus periods, flanking a period of blue light illumination. The haemodynamic response to optogenetic stimulation in the Thy1-ChR2-EYFP mouse model was previously observed to last about 45 seconds [30], similar to data collected in other optogenetic models of mural cell stimulation where arteries were observed to dilate for approximately 60 and 20 seconds respectively, and capillaries were observed to dilate on the scale of minutes [6], [23]. We thus constrained our volumetric acquisitions to 45 seconds. Each volumetric scan lasted for 42.98 seconds and used the resonant scanner with 5x frame averaging (resulting in a signal-to-noise ratio on the vasculature channel of 6.9±1.6), a nominal lateral resolution of 0.99 μm, a z-step of 2.64 μm, and traversed 250.8 μm of cortical depth. This z-step was chosen because larger z-steps led to discontinuities in the capillary bed segmentation. We imaged the brain in sections from the cortical surface to a depth of 250 μm and from 250 μm to 500 μm of cortical depth. The acquisitions were split into two slabs to maximize the volume coverage without creating discontinuities in the capillary segmentation masks while maintaining a scan duration of under 45 seconds. Stacks were acquired either from the cortical surface down or from the bottom of the stacks toward the cortical surface on different paired acquisitions, so that the post-stimulus delay of image acquisition at different cortical depths was not constant. A 239-um diameter circular photostimulation region was positioned 250 μm beneath the cortical surface. The photostimulation light was raster scanned over this ROI for 5 seconds between imaging scans, with a pixel dwell time of 4 μs for each pixel of the stimulation ROI. Each pixel in the photostimulation ROI was thus excited every 501.55 milliseconds, i.e. at about 2Hz. We used two laser powers for 458-nm photostimulation: 1.1 mW/mm2 and 4.3 mW/mm2. The 1.1 mW/mm2 is just above the threshold for ChR2 photoactivation and is expected to elicit a small degree of neuronal activity, whereas the 4.3 mW/mm2 photoactivation elicits more neuronal activity while still being below ChR2 saturation level [33], [36]. Low powers were used to activate ChR2-expressing neurons to minimize tissue heating [37]. We also performed control experiments via 552-nm photostimulation (where ChR2 excitation falls to 4% of its peak) at 4.3 mW/mm2 [38]. Mice rested for approximately 5-7 minutes between scan pairs for the vascular tone to return to its resting state. Photostimulation parameters were presented in randomized order for each mouse.

Revised version: Functional scans were acquired over baseline and post-stimulus periods, flanking a period of blue light illumination. The haemodynamic response to optogenetic stimulation in the Thy1-ChR2-EYFP mouse model was previously observed to last about 45 seconds, and observed to last for minutes in our experimental model (cf. Supplementary Figure 2) [30] These vascular radius time courses were similar to data collected in other optogenetic models of mural cell stimulation where arteries were observed to dilate for approximately 60 and 20 seconds respectively, and capillaries were observed to dilate on the scale of minutes [6], [23]. We thus constrained our volumetric acquisitions to 45 seconds. Each volumetric scan lasted for 42.98 seconds and used the resonant scanner with 5x frame averaging (resulting in a signal-to-noise ratio on the vasculature channel of 6.9±1.6), a nominal lateral resolution of 0.99 μm, a z-step of 2.64 μm, and traversed 250.8 μm of cortical depth. This z-step was chosen because larger z-steps led to discontinuities in the capillary bed segmentation. We imaged the brain in sections from the cortical surface to a depth of 250 μm and from 250 μm to 500 μm of cortical depth. The acquisitions were split into two slabs to maximize the volume coverage without creating discontinuities in the capillary segmentation masks while maintaining a scan duration of under 45 seconds. Stacks were acquired either from the cortical surface down or from the bottom of the stacks toward the cortical surface on different paired acquisitions, so that the post-stimulus delay of image acquisition at different cortical depths was not constant. A 239-μm diameter circular photostimulation region was positioned 250 μm beneath the cortical surface. The photostimulation light was raster scanned over this ROI for 5 seconds between imaging scans, with a pixel dwell time of 4 μs for each pixel of the stimulation ROI. Each pixel in the photostimulation ROI was thus excited every 501.55 milliseconds, i.e. at about 2Hz. We used two laser powers for 458-nm photostimulation: 1.1 mW/mm2 and 4.3 mW/mm2. The 1.1 mW/mm2 is just above the threshold for ChR2 photoactivation and is expected to elicit a small degree of neuronal activity, whereas the 4.3 mW/mm2 photoactivation elicits more neuronal activity while still being below ChR2 saturation level [33], [36]. Low powers were used to activate ChR2-expressing neurons to minimize tissue heating [37]. We also performed control experiments via 552-nm photostimulation (where ChR2 excitation falls to 4% of its peak) at 4.3 mW/mm2 [38]. Mice rested for approximately 5-7 minutes between scan pairs for the vascular tone to return to its resting state. Photostimulation parameters were presented in randomized order for each mouse.

Segmentation and Graph Extraction

Ground truth generation

To generate ground truths, we used ilastik’s pixel classification workflow, which relies on random forests, to annotate blood vessels and pyramidal excitatory neurons soma’s in 42 volumes from 25 Thy1-ChR2-EYFP mice for training our DL model (24 images from 16 Thy1-ChR2-EYFP mice from other studies were included in the training, validation, and test cohorts to increase our dataset size). Images were semi-automatically segmented in groups of up to 4 each, with a size of 507×507×250 μm each, due to ilastik’s inability to handle large amounts of annotations over more than a few volumes. The rater labelled targets (neurons vs. blood vessels), corrected mistakes, and classified the pixels where the output was uncertain (as shown using ilastik’s uncertainty guidance feature). During the manual annotation, feature selection was repeatedly optimized using ilastik’s suggest feature function (wrapper method), leading to different features being used for the random forest model for each set of images. Small (under 50 pixels) isolated vascular components were removed with connected component analysis in Python using scikit-image and connected-components-3d. Testing data was withheld until the final model was selected based on the models’ performance evaluation on the validation data set.

Data preprocessing and augmentation

We used the MONAI python package for data preprocessing, augmentation, model training, and prediction [39]. All images and ground truth segmentation masks were up-sampled to an isotropic voxel size using bilinear and nearest neighbour interpolation. Raw 10-bit image intensities were normalized to range from 0 to 1.0. Each volume had eight 128×128×128 pixel patches randomly cropped out for training. Data augmentation, transformations and parameters are listed in Supplementary Table 2. Spatial transformations were selected to expand data variety via cropping, rotations, and mirroring hence exposing the network to images that would be acquired on different positioning of the animal under the microscope. Zooming and deformation transformations were included to expose the network to small changes in the size and morphology of the vasculature. Intensity and Gaussian transformations exposed the network to signal intensity and contrast variations. Dropping pixels was included as the resonant scanner acquisition occasionally yields images with some relatively low signal pixels.

Model Architecture

Earlier version: We trained a state-of-the-art 3D vision transformer (UNETR) model and the U-Net model for baseline comparison, as implemented in PyTorch using the MONAI library [39]. UNETR is a U-Net style architecture with a transformer-based multi-attention head encoder and a CNN decoder. [40], [41]. The encoder takes an input image (or patch from each batch) and breaks it down into a sequence of non-overlapping patches, each of size 16×16×16 pixels, which are weighted differently to account for variation in signal intensities and patterns within an image. The sequence of patches is then passed through a multi-head self-attention and multilayer perceptron encoder to capture self-attention between different pixels of the patch to encode long-range relationships between patches. The encoder is then connected to a CNN decoder with skipped connections to the encoder to map features back onto the original image at multiple spatial scales [40]. The model was trained on 2-channel images (EYFP-expressing neurons and Texas Red-labelled vasculature) as inputs to segment neurons and vasculature. We chose the following parameters for the UNETR architecture: a 12 multi-attention head encoder, 16 convolutional features in the first layer of the encoder, a hidden layer size of 768, a multilayer perceptron size of 3072, and Monte Carlo dropout, where on each run of the model 10% of weights were zeroed [40]–[42].

Revised version: We trained a state-of-the-art 3D vision transformer (UNETR) model and the U-Net model for baseline comparison, as implemented in PyTorch using the MONAI library [39]. UNETR is a U-Net style architecture with a transformer-based multi-attention head encoder and a CNN decoder [40], [41]. The encoder takes an input image (or patch from each batch) and breaks it down into a sequence of non-overlapping patches, each of size 16×16×16 pixels, which are weighted differently to account for variation in signal intensities and patterns within an image. The sequence of patches is then passed through a multi-head self-attention and multilayer perceptron encoder to capture self-attention between different pixels of the patch to encode long-range relationships between patches. The encoder is then connected to a CNN decoder with skipped connections to the encoder to map features back onto the original image at multiple spatial scales [40]. The model was trained on 2-channel images (EYFP-expressing neurons and Texas Red-labelled vasculature) as inputs to segment neurons and vasculature. We chose the following parameters for the UNETR architecture: a 12 multi-attention head encoder, 16 convolutional features in the first layer of the encoder, a hidden layer size of 768, a multilayer perceptron size of 3072, and Monte Carlo dropout, where on each run of the model 10% of weights were zeroed [40], [41], [42].

Earlier version: The U-net model used is based on the original 2D U-net proposed by Ronneberger et al. in 2015 and extended into 3D via 3D convolutional operations and features residual blocks, parametric rectifying linear units, and instance normalization as described by Kerfoot et al. [41], [43]–[47]. The residual blocks were implemented for better gradient flow during the network training. At the same time, the parametric ReLU activation functions were used to improve the ability of the model to adjust its weights during training. Instance normalization was implemented to reduce contrast differences between images fed into the network to improve model robustness. The model was implemented with five layers featuring 16, 32, 64, 128, and 256 channels and dropout. The U-net model was trained using the same data augmentation employed during the UNETR model training.

Revised version: The U-net model used is based on the original 2D U-net proposed by Ronneberger et al. in 2015 and extended into 3D via 3D convolutional operations and features residual blocks, parametric rectifying linear units, and instance normalization as described by Kerfoot et al. [41], [43], [44], [45], [46], [47]. The residual blocks were implemented for better gradient flow during the network training. At the same time, the parametric ReLU activation functions were used to improve the ability of the model to adjust its weights during training. Instance normalization was implemented to reduce contrast differences between images fed into the network to improve model robustness. The model was implemented with five layers featuring 16, 32, 64, 128, and 256 channels and dropout. The U-net model was trained using the same data augmentation employed during the UNETR model training.

Model Hyperparameter Optimization, Training, and Prediction

We used a grid search to determine optimal hyperparameters for the UNETR and UNet models within the following parameter space [43], [48]: loss functions (Dice Loss, Dice + Cross Entropy Loss, Dice + Focal Loss, or Tversky Loss), dropout rates (0.1, 0.2, or 0.3), learning rates (5e-3, 1e-3, 5e-4, 1e-4, or 1e-5), and the number of residual units for the Unet model (2 or 3) [49], [50]. We utilized the Adam optimizer [51], and the best model during hyperparameter optimization was selected based on the validation Dice similarity coefficient (DSC) and trained for a maximum of 2400 epochs. The Dice score was the principal evaluation metric, to maximize the overlap between ground truth and prediction masks. Precision and recall were used as secondary metrics to achieve balanced precision and recall where over-segmented results were produced. For early stopping, the epoch with the best performance of the DSC on the validation dataset was selected. The final model that had the best Dice score during hyper-parameter optimization was trained with a Dice + Cross Entropy Loss function, a dropout rate of 0.1, 2 residual units, a learning rate of 1e-5, and a batch size of 1 image with 8 crops per image. Training and optimization were performed on the Narval cluster of Calcul Quebec and the Digital Research Alliance of Canada, with each node using 498 GB of RAM and 4 Nvidia A100 GPUs, each with 40GB HBM2 VRAM.

Ilastik utilized a random forest model from the Vigra library with the default parameters and 100 trees [52], [53]. We initially utilized all default 3D Colour/Intensity/Texture/Edge features during ilastik feature selection and added 2D Colour/Intensity/Texture/Edge features with a sigma of 20.0 pixels. These 2D features were added as a strong, largely planar artifact was observed towards the surface of the images in the neuron channel. We used the wrapper method for feature selection with a set size penalty of 0.10 to determine the optimal features from the starting set. The model was then trained using live updates on the Narval cluster of Calcul Quebec and the Digital Research Alliance of Canada, with each node using 249 GB of RAM and 2 x AMD Rome 7532 CPUs with 64 cores.

For the ilastik model, a subset of 3 training images from 3 mice was used as ilastik’s random forest was unable to train a model with more data even on compute nodes with large memory and CPU resources. The model could never complete training within the time limits of the resource allocations. This inability to complete model optimization was expected as ilastik’s documentation does not recommend training random forest models with full annotation datasets, and increasing the number of annotations does not necessarily lead to better predictions for the pixel classification workflow [52]. The ilastik random forest model was trained with whole images rather than with patches.

Model Comparison

Earlier version: To compare models, we employed several metrics evaluating the similarity between the ground truth and prediction: the Dice Similarity Coefficient (DSC), Precision, and Recall, as well as surface-based metrics: 95% Hausdorff distance (HD95) and mean surface distance. We specifically focused on mean surface distance and HD95 distance since the centerline extraction used in the downstream analysis was highly sensitive to surface irregularities. The model with the lowest standard deviation on the mean surface distance was to be selected absent statistically significant differences in the average mean surface distance between different models. Consistency in surface placement was key for extracting good centerlines and graphs from the segmentation masks. The metrics utilized are defined as follows:

Revised version: To compare models, we employed several metrics evaluating the similarity between the ground truth and prediction: the Dice Similarity Coefficient (DSC), Precision, and Recall, as well as surface-based metrics: 95% Hausdorff distance (HD95) and mean surface distance. We specifically focused on mean surface distance and HD95 distance since the centerline extraction used in the downstream analysis was highly sensitive to surface irregularities. The model with the lowest standard deviation on the mean surface distance was to be selected absent statistically significant differences in the average mean surface distance between different models.

Added in the revised version: Consistency in surface placement was key for extracting good centerlines and graphs from the segmentation masks. The metrics utilized are defined as follows:

Earlier version: The DSC measures the overlap between the prediction and ground truth with a value of 1 corresponding to complete overlap. Precision measures the rate of correctly returned predicted values, whereas recall accesses the rate of return of targeted values. Hausdorff 95% distance (HD95) provides the maximum of the 95th percentile of the minimum surface distance between the boundaries of the ground truth and the predicted segmentation mask. The mean surface distance measures the average minimum distance between the surface of the ground truth and the surface of the predicted segmentation mask. Together, the HD95 and mean surface distance assess model performance at the boundaries of the segmentation masks.

Revised version: The DSC measures the overlap between the prediction and ground truth with a value of 1 corresponding to complete overlap, X represents the ground truth and Y represents the predicted value. Precision measures the rate of correctly returned predicted values, whereas recall accesses the rate of return of targeted values. For the Hausdorff 95% distance (HD95), the distance d is the infimum of the Euclidean distance between points x and y (from the surface of segmentation masks X and Y) to segmentation masks Y and X, respectively. HD95 computes the maximum of the 95th percentile of the supremums (sup) of these minimum surface distances (d(x,Y) and d(X,y)) between the boundaries of the ground truth and the predicted segmentation mask. The mean surface distance measures the average minimum distance between the boundary of the ground truth and the boundary of the predicted segmentation mask. In Equation (5), nx and ny are the sets of boundary points of X and Y, respectively. The infimum of the minimum Euclidean distance, d, from boundary points p from nx, and q from ny to sets nx and ny are then computed and averaged to compute the Mean Surface Distance. Together, the HD95 and Mean Surface Distance assess model performance at the boundaries of the segmentation masks.

Graph extraction of cerebrovascular networks

Earlier version: Graph extraction was performed in Python except for the centerline calculation, which was performed in MatLab (R2021a). Upsampled image acquisitions (0.99 μm isotropic voxel size) were registered with ANTsPy using the rigid registration method with a total sigma of 2 pixels (for smoothing within the registration function) and mean squared error as the similarity metric as input parameters of the registration function (Figure 2A,B). We selected the first baseline scan from each region of interest that was scanned from the bottom to the top to serve as a reference to which all other images from the same region were aligned. The calculated transforms were used to transform images from the same ROI to the space of the reference image using linear interpolation [54]. We then generated segmentation masks for each aligned image using our trained UNETR model with sliding window inference, and we retained Monte Carlo dropout during prediction to create an ensemble of 20 UNETR models, so as to assess the model uncertainty (Figure 2D) [55]. To extract the centerlines, we assumed that no background was present within the vessels, as regions of background within vessels disrupt the centerline extraction algorithm. Background labelled pixels surrounded by blood vessel segmentations were thus filled in using connected component analysis. Disconnected vascular components under 50 pixels were assumed to be noise and removed. Each aligned binary segmentation mask was dilated three times with a footprint defined by scikit-image’s disk function with a radius of 1 pixel. Next, we computed the union of all segmentation masks from the same ROI generated from registered images (Figure 2E). This common segmentation mask for all time points was eroded to a centerline via a medial axis transformation (bwskel function in Matlab (R2021a)). Critically, using the union of all time points minimized the sensitivity of centreline identification to RBC stalls at individual time points. “Hair” segments shorter than 20 micrometres and terminal on one end were removed to reduce false positive vascular branches [56]. The vascular centerlines were next used to construct vascular graphs (with sknw) for each ROI, rendering coordinates of vertices along the vessel’s centreline as edges, and branch points as nodes [57].

Revised version: Graph extraction was performed in Python except for the centerline calculation, which was performed in MatLab (R2021a). Upsampled image acquisitions (0.99 μm isotropic voxel size) were registered with ANTsPy using the rigid registration method with a total sigma of 2 pixels (for smoothing within the registration function) and mean squared error as the similarity metric as input parameters of the registration function (Figure 2A,B). We selected the first baseline scan from each region of interest that was scanned from the bottom to the top to serve as a reference to which all other images from the same region were aligned. The calculated transforms were used to transform images from the same ROI to the space of the reference image using linear interpolation [54]. We then generated segmentation masks for each aligned image using our trained UNETR model with sliding window inference, and we retained Monte Carlo dropout during prediction to create an ensemble of 20 UNETR models, so as to assess the model uncertainty (Figure 2D) [55].

Computational analysis pipeline:

A. The stacks of 2PFM slices were registered using ANTS rigid registration and aligned to the reference time point. B. Images were upsampled using bicubic interpolation to an isotropic resolution of 0.99×0.99×0.99 μm. C. An ensemble of UNETR deep learning models with dropout generated segmentation masks at each time point, producing probability maps. D. The mean and standard deviation of the probability of each pixel being vasculature were computed and used to create binary vascular segmentation masks. E. The union over the vascular segmentation masks for all time points was computed, and background pixel clusters within vessel masks were removed. F. The vascular segmentation mask was thinned down to centerlines and rendered as a graph, where edges were vessel segments connecting branch points (nodes). This skeleton was overlaid on the vasculature channel from which the neuron channel was subtracted. G. The plane orthogonal to the tangent to the vessel’s travel direction was computed every micrometre along the centerline. H,I. 1D signal intensity profiles at each centerline vertex were computed in the orthogonal plane every 10°. J. The boundary for each profile was placed at the minimum of the signal gradient for that signal intensity profile. K. The raw intensity image with the detected boundary points, where outlier boundary points (in green) were defined as points over 2 standard deviations from the mean and excluded. L. Visualization of the changes in vertex-wise radii on a sample vascular network.

Added in the revised version: To extract the centerlines, we assumed that no background was present within the vessels, as regions of background within vessels disrupt the centerline extraction algorithm. Background labelled pixels surrounded by blood vessel segmentations were thus filled in using connected component analysis. Disconnected vascular components under 50 pixels were assumed to be noise and removed. Each aligned binary segmentation mask was dilated three times with a footprint defined by scikit-image’s disk function with a radius of 1 pixel. Next, we computed the union of all segmentation masks from the same ROI generated from registered images (Figure 2E). This common segmentation mask for all time points was eroded to a centerline via a medial axis transformation (bwskel function in Matlab (R2021a)). Critically, using the union of all time points minimized the sensitivity of centreline identification to red blood cell (RBC) stalls at individual time points. “Hair” segments shorter than 20 μm and terminal on one end were removed to reduce false positive vascular branches: this functionality of the skeletonization algorithm is currently unavailable in any python implementation, but is available in Matlab [56]. The vascular centerlines were next used to construct vascular graphs (with sknw) for each ROI, rendering coordinates of vertices along the vessel’s centreline as edges, and branch points as nodes [57].

Based on the vascular graphs, we computed vascular radius estimates at each vertex at each timepoint and then calculated the vertex’s distance to the nearest YFP-expressing pyramidal neuron (as measured by the distance transformation at the vertex to the nearest labelled neuron, with the radius of the vessel subtracted off). We employed a two-stage approach to estimate the radius for each point on the centerline. First, using the binary segmentation mask, we calculated the distance from every vessel pixel in the mask to the nearest background pixel. These values were averaged for each vessel to estimate its radius. The radius estimate defined the size of the Gaussian kernel that was convolved with the image to smooth the vessel: smaller vessels were thus convolved with narrower kernels.

Earlier version: In the second stage, the registered raw image was deconvolved with the point spread function of the microscope, as measured via FluoSpheres carboxylate-modified Microspheres (Cat# F8803, Thermo Fisher Scientific Inc, Waltham MA), using Richardson-Lucy deconvolution. To low pass filter the centreline in advance of computing the tangent vector at each vertex, the coordinates of the vertices along the centerline were smoothed using a Gaussian with a sigma of 3. Next, the tangent vector to the centerline was specified by calculating the gradient of the centerline path. Given this local tangent to the vessel’s direction of travel at a given vertex, we extracted image intensities in the orthogonal plane from the deconvolved raw registered image. A 2D Gaussian kernel with sigma equal to 80% of the estimated vessel-wise radius was used to low-pass filter the extracted orthogonal plane image and find the local signal intensity maximum searching, in 2D, from the center of the image to the radius of 10 pixels from the center. The orthogonal plane image was sampled every 10 degrees (as finer radial sampling did not improve the estimation) along radial lines emanating from the local signal intensity maximum closest to the center of the image and 5x bicubic upsampled to extract thirty-six 1D signal intensity profiles at that vertex, as shown in Figure 2H. These line profiles were then convolved with a 1D Gaussian kernel with a sigma of 80% of the estimated radius of that vessel, and the gradient of each profile was calculated as shown in Figure 2I, J. The vessel boundary was then placed at the local minimum of the gradient of each profile. The mean and standard deviation of the boundary distances for the thirty-six 1D line profiles were calculated, and boundary points greater than 2 standard deviations away from the mean were excluded (Figure 2K). This radius estimation procedure was repeated for all vertices of all vessels.

Revised version: In the second stage, the registered raw image was deconvolved with the point spread function of the microscope, as measured via FluoSpheres carboxylate-modified Microspheres (Cat# F8803, Thermo Fisher Scientific Inc, Waltham MA), using Richardson-Lucy deconvolution. To low pass filter the centreline in advance of computing the tangent vector at each vertex, the coordinates of the vertices along the centerline were smoothed using a Gaussian with a sigma of 3. Next, the tangent vector to the centerline was specified by calculating the gradient of the centerline path. Given this local tangent to the vessel’s direction of travel at a given vertex, we extracted image intensities in the orthogonal plane from the deconvolved raw registered image. A 2D Gaussian kernel with sigma equal to 80% of the estimated vessel-wise radius was used to low-pass filter the extracted orthogonal plane image and find the local signal intensity maximum searching, in 2D, from the center of the image to the radius of 10 pixels from the center. The orthogonal plane image was sampled every 10 degrees (as finer radial sampling did not improve the estimation cf. Supplementary Figure 3) along radial lines emanating from the local signal intensity maximum closest to the center of the image and 5x bicubic upsampled to extract thirty-six 1D signal intensity profiles at that vertex, as shown in Figure 2H. These line profiles were then convolved with a 1D Gaussian kernel with a sigma of 80% of the estimated radius of that vessel, and the gradient of each profile was calculated as shown in Figure 2I, J. The vessel boundary was then placed at the local minimum of the gradient of each profile. The mean and standard deviation of the boundary distances for the thirty-six 1D line profiles were calculated, and boundary points greater than 2 standard deviations away from the mean were excluded (Figure 2K). This radius estimation procedure was repeated for all vertices of all vessels.

In the last step, we computed the distance from each image voxel to the nearest YFP-expressing pyramidal neuron by computing the distance from every image pixel not belonging to the neuron mask to the closest YFP-expressing neuron’s soma boundary (based on the neuron’s binary segmentation mask). At the vessel vertices, these distances were adjusted by subtracting the vessel’s radius. We thereby captured the distance from the vessel surface to the closest YFP-expressing neuron.

Boundary detection validation

To assess the uncertainty in the radius estimates at each vertex, we simulated changes to the vascular diameter by “resizing” the extracted orthogonal plane image and adding Gaussian noise to it. These steps were undertaken to validate the ability of the boundary detection method to estimate diameter changes and evaluate the robustness of the estimates to increased amounts of noise. For the diameter change estimation, images were resampled from the orthogonal plane by a random factor, uniformly distributed in the range of 0.5x to 2x. Then, the aforementioned boundary detection methods were used to estimate vertex-wise calibre changes. The end goal of the assessment was to see how close the boundary detection method-based radius change matched the prescribed diameter change. In the second task, we added Gaussian noise with a sigma randomly chosen from a uniform distribution ranging from 0 to 500 SU to the orthogonal plane image. The noise was added to the image after deconvolution with the PSF and the extraction of the orthogonal plane but before Gaussian smoothing. We then reported on the percent change in the radius as a result of resizing or adding noise in relation to the baseline radius.

Vascular Network Analysis

Earlier version: Leveraging the graph representation of the vasculature in our pipeline, we next used graph theory to better understand the networks’ behaviour upon neuronal activation. We looked at morphometrics including vascular segment count density, vascular length density, and mean vessel length for comparison with other work. To demonstrate the benefits of extracting a graphical representation of the vasculature in situ, we looked at graph theory metrics including the assortativity of vascular radius changes and changes to the global efficiency of the capillary (below 5 μm in radius) network following optogenetic stimulation. The assortativity measures the tendency for one vessel to dilate or constrict by a similar amount as its neighbours. The assortativity (q) of radius changes in response to stimulation is defined as the Pearson correlation coefficient of these changes on connected vessels: where exy represents the fraction of vessels (edges) in the network that join together nodes with values x and y (i.e. radius changes with values x and y); ax and by are the percentages of edges connecting nodes with values x and y, and σa and σb are the standard deviations of ax and by [58]. The efficiency, in turn, captures how easily the graph can be traversed. For vascular graphs, efficiency can be conceptualized as the average of the inverse of the total resistive distance of the shortest paths of all combinations of vascular junctions, with the hydraulic resistivity serving as the distance between them. The resistivity (equation 7) is summed along the shortest paths, and the inverse of this sum is then averaged across all shortest paths to compute the efficiency (equation 8). Finally, pPhotostimulation-induced change in the efficiency, from its baseline level, is reported.

Revised version: Leveraging the graph representation of the vasculature in our pipeline, we next used graph theory to better understand the networks’ behaviour upon neuronal activation. We looked at morphometrics including vascular segment count density, vascular length density, and mean vessel length for comparison with other work. To demonstrate the benefits of extracting a graphical representation of the vasculature in situ, we looked at graph theory metrics including the assortativity of vascular radius changes and changes to the global efficiency of the capillary (below 5 μm in radius) network following optogenetic stimulation. The assortativity measures the tendency for one vessel to dilate or constrict by a similar amount as its neighbours. The assortativity (q) of radius changes in response to stimulation is defined as the Pearson correlation coefficient of these changes on connected vessels: where exy represents the fraction of vessels (edges) in the network that join together nodes with values x and y (i.e. radius changes with values x and y); ax and by are the percentages of edges connecting nodes with values x and y, and σa and σb are the standard deviations of ax and by [58]. The efficiency, in turn, captures how easily the graph can be traversed. For vascular graphs, efficiency can be conceptualized as the average of the inverse of the total resistive distance of the shortest paths of all combinations of vascular junctions, with the hydraulic resistivity serving as the distance between them. The resistivity (equation 7) is summed along the shortest paths, and the inverse of this sum is then averaged across all shortest paths to compute the efficiency (equation 8). Finally, photostimulation-induced change in the efficiency, from its baseline level, is reported.

In computing resistivity, we assumed a fluidic viscosity (μ) of 4 cP which is within the physiological range [59]; L was the length of the capillary; and R was the radius of the capillary. Vessels greater than 10 μm in diameter were excluded from the efficiency calculation as we wanted to examine blood flow through the capillary nexus between arteries and veins where blood flow may be reversible [60]. The sum of the resistivities along the least resistive path specified ⍴ij in the efficiency calculation: this sum quantifies how hard it is for blood to move through the said microvascular bed. N refers to the number of nodes in the network. The efficiency was calculated as the average of the inverse of the least resistive paths between all pairs of nodes. The efficiency increases with increasing radius on the shortest paths through the network.

Statistical models

Earlier version: To statistically compare deep learning model performance, we used the Wilcoxon signed-rank test as implemented in scipy (1.9.3) [61]. Statistics for vascular radius and network metrics were performed in R (4.3.1) using restricted maximum likelihood mixed effects models as implemented in lme4 (1.134) [62], with post hoc comparisons done using emmeans (1.8.8) [63]. We ran the following linear mixed effects model, separately on dilating and constricting vessels that exhibited radius changes above two standard deviations of the vessel’s baseline radius, to examine the effects of photostimulation power on microvascular radius changes in responding vessels:

Revised version: To statistically compare deep learning model performance, we used the Wilcoxon signed-rank test as implemented in scipy (1.9.3) [61]. Statistics for vascular radius and network metrics were performed in R (4.3.1) using restricted maximum likelihood mixed effects models as implemented in lme4 (1.134) [62], with post hoc comparisons done using emmeans (1.8.8) [63]. We ran the following linear mixed effects model, separately on dilating and constricting vessels that exhibited radius changes above two standard deviations of the vessel’s baseline radius, to examine the effects of photostimulation power on microvascular radius changes in responding vessels: Due to the difference in their vessel wall ultrastructure, larger microvessels (above 5 μm in radius) were examined separately from capillaries (radius < 5 μm).

Removed in the revised version: ΔRadius ∼ Stimulation + (1|Vessel)

Removed in the revised version: Due to the difference in their vessel wall ultrastructure, larger microvessels (above 5 μm in radius) were examined separately from capillaries (radius < 5 μm).

Earlier version: For graph metrics, we included nesting of the field of view within each subject as a random effect so as to account for differences in vascular network architecture within an individual. The linear mixed effects models for the graph metrics used were as follows:

Revised version: For graph metrics, we included nesting of the field of view within each subject as a random effect so as to account for differences in vascular network architecture within an individual. The linear mixed effects models for the graph metrics used were as follows: In Tables and text, all values have been quoted as mean ± standard deviation, unless otherwise specified.

Removed in the revised version: Assortativity ∼ Stimulation + (1|Subject/Field of View)

Removed in the revised version: Δ Efficiency ∼ Stimulation + (1|Subject/Field of View)

Removed in the revised version: In Tables and text, all values have been quoted as mean ± standard deviation, unless otherwise specified.

Results

Application of our computational pipeline resulted in robust segmentation of the vasculature and neurons from 4D in situ 2PFM images and rendering of the microvasculature as a graph. The vessel-wise and vertex-wise calibres were tracked across stimulation conditions and related to the cortical depth and the distance to the closest YFP-expressing neuron, mapping the network-level vascular responses to ChR2 activation and revealing the coordination of the microvascular responses following neuronal activation.

Segmentation model results and comparisons

Earlier version: We compared an ensemble of UNETR models, an ensemble of U-Net models, and an ilastik random forest model on a test dataset of 9 images (507×507×250 μm each) from 6 mice. Examples of segmentation masks produced by each of the models are shown in Figure 4. When evaluating model performance, we paid close attention to the smoothness of the surface of the segmentation masks due to the sensitivity of the centerline extraction algorithms to irregularities in the surface of the masks: smoother vascular segmentation masks resulted in fewer falsely identified vessel branches. Ilastik tended to over-segment vessels, i.e. the model returned numerous false positives, having a high recall (0.89±0.19) but low precision (0.37±0.33) (Figure 3, Supplemental Table 3). When comparing the UNETR and U-Net models, we focused on the surface-based mean surface distance and HD95 distance metrics. Since we observed no significant differences in these metrics between the two models, we selected the UNETR model as our final model because it produced more consistent segmentations on visual inspection and showed significantly better performance than ilastik on HD95 for both vessel and neuron segmentation (p

Revised version: We compared an ensemble of UNETR models, an ensemble of U-Net models, and an ilastik random forest model on a test dataset of 9 images (507×507×250 μm each) from 6 mice. Examples of segmentation masks produced by each of the models are shown in Figure 4. When evaluating model performance, we paid close attention to the smoothness of the surface of the segmentation masks due to the sensitivity of the centerline extraction algorithms to irregularities in the surface of the masks: smoother vascular segmentation masks resulted in fewer falsely identified vessel branches. Ilastik tended to over-segment vessels, i.e. the model returned numerous false positives, having a high recall (0.89±0.19) but low precision (0.37±0.33) (Figure 3, Supplementary Table 3). When comparing the UNETR and U-Net models, we focused on the surface-based mean surface distance and HD95 distance metrics. Since we observed no significant differences in these metrics between the two models, we selected the UNETR model as our final model because it produced more consistent segmentations on visual inspection and showed significantly better performance than ilastik on HD95 for both vessel and neuron segmentation (p < 0.05).

Model performance metrics:

The Dice, precision, recall, mean surface distance, and HD95 distance for the vascular (A) and neuron (B) channels. Each model was evaluated on the same test dataset composed of 9 images (250×507×507 μm each) from 6 mice. A Wilcoxon signed-rank test was used to compare the model’s performance on each performance metric for images from the test dataset. * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted.

Visual model comparison:

A. Raw images of the vascular channel with the neuron channel subtracted to facilitate vessel visualization. The first and last stacks in each row span from the cortical surface to 250 μm below the surface, while the middle stack spans from 250 μm below the surface to 500 μm below the surface. All images were from the test dataset, which was unseen during model training. B. Ground truth segmentation masks for the vasculature were generated by a rater who utilized ilastik-assisted manual segmentation. C. Ilastik predictions generated via a random forest model. D. Binary segmentation masks generated by an ensemble of 3D UNet Models. E. Binary segmentation masks generated by an ensemble of 3D UNETR models.

Removed in the revised version: < 0.05).

Vessel extraction improvements via image registration

Rigid registration across all time points from the same field of view improved the ability to trace vascular paths from in situ 2PFM data. Firstly, registration decreased the mean squared error (MSE) between acquisitions from 1306±747 to 0.008±0.003 Signal Units. The number of images acquired per field of view ranged from a total of 2 to 10 depending on how many repeats were able to be acquired. Registration enabled the computation of the union of segmentation masks from all time points. This increased the number of vessel segments identified in each field of view from 241±174 based on a single time point to 412±281 vessel segments per field of view (507×507×250 μm, n=107 fields of view). Taking the union of segmentation masks of an image stack across all time points substantially decreased the incidence of gaps in capillaries, likely arising due to “transient RBC plugs.” The pipeline’s ability to reconstruct the cortical vascular network was thus significantly improved by registering data obtained at different time points.

Validation of pipeline sensitivity to geometric changes

To evaluate the ability of our computational pipeline to detect vessel caliber changes of various magnitudes, we simulated a range of vascular caliber changes and injected various levels of Gaussian noise. Across >100,000 simulations, the fit of the estimated radius following rescaling against the simulated radius had an R2 value of 0.68. Figure 5B presents a heatmap of the estimated radius post-scaling vs. simulated radius, across different vertices of vessel centerlines, highlighting the ability of our pipeline to estimate vascular radii accurately. The addition of Gaussian noise revealed the robustness of the pipeline: radius estimates remained stable with increasing noise levels, until the addition of noise with a standard deviation of over 200 SU (with the intensity of the images ranging from 0 to 1023 SU).

Estimation of simulated radii changes:

A. An image in the plane orthogonal to the local tangent to a capillary with the detected boundary (in blue) and with the estimated radius of 2.28 μm. On the right, this image was resized (upsampling, via bicubic interpolation, by 1.10 times) to simulate dilation. B. The plot shows correspondence between the estimated radius following scaling and the simulated level of scaling. C. An image in the plane orthogonal to the local tangent of a capillary with the detected boundary (in blue) and with the estimated radius of 3.65 μm. On the right, Gaussian noise with a sigma of 205.36 SU was added to the image. D. The estimated % change in the vessel’s radius after the addition of varying levels of Gaussian noise, demonstrating the robustness of the radius estimated to noise.

Vascular morphology and heterogeneity within and among vessels

Earlier version: Segmentation coupled with graph extraction enabled a detailed characterization of the microvascular network properties. The morphological properties of extracted networks are listed in Table 1, with the probability densities of the vessel length, baseline vessel radius, mean vessel segment depth, and vessel branch point depth shown in Supplementary Figure 3. On the extracted graphs, the vascular radius and distance to labelled neurons were sampled every 1-1.73 μm, enabling detailed analysis of the relationship between the vessel radius change and the proximity to the YFP-expressing neurons. The radius was tracked across different time points, permitting the analysis of the stimulation-induced change in the vascular caliber. To highlight the ability of the pipeline to detect vessels that significantly change their radius after stimulation, Figure 6A shows the standard deviation of the average radii on each vessel segment during baseline frames for three mice. There was a large difference in this standard deviation across various blood vessels, showcasing the model’s ability to reveal baseline variations within each subject. We examined the average change in the vascular radius of each vessel segment after vs. before photostimulation (Figure 6B), with even finer spatial patterns detected by analyzing the vertex-wise radius changes (Figure 6C). The vascular diameter changes were related to the distance from the vessel’s surface to the closest pyramidal neuron at each vertex of the centerline (Figure 6D). The vertex-wise radii estimation allowed the assessment of the variations in radii changes within individual blood vessels (Figure 7). Notably, capillary radius varied along the vessel length across the baseline frames, by 24 ± 28% of the mean resting radius. Consequently, point measurements in vessel calibers - that are widely reported in the literature - do not permit accurate estimation of the microvessel volume changes. Together, the within- and across-vessel radii estimations illustrate the pipeline’s ability to capture spatial variations in the vascular reactivity and relate it to other morphological features (e.g., the distance to the closest activated neuron).

Revised version: Segmentation coupled with graph extraction enabled a detailed characterization of the microvascular network properties. The morphological properties of extracted networks are listed in Table 1, with the probability densities of the vessel length, baseline vessel radius, mean vessel segment depth, and vessel branch point depth shown in Supplementary Figure 4. On the extracted graphs, the vascular radius and distance to labelled neurons were sampled every 1-1.73 μm, enabling detailed analysis of the relationship between the vessel radius change and the proximity to the YFP-expressing neurons. The radius was tracked across different time points, permitting the analysis of the stimulation-induced change in the vascular caliber. To highlight the ability of the pipeline to detect vessels that significantly change their radius after stimulation, Figure 6A shows the standard deviation of the average radii on each vessel segment during baseline frames for three mice. There was a large difference in this standard deviation across various blood vessels, showcasing the model’s ability to reveal baseline variations within each subject. We examined the average change in the vascular radius of each vessel segment after vs. before photostimulation (Figure 6B), with even finer spatial patterns detected by analyzing the vertex-wise radius changes (Figure 6C). The vascular diameter changes were related to the distance from the vessel’s surface to the closest labelled pyramidal neuron at each vertex of the centerline (Figure 6D). The vertex-wise radii estimation allowed the assessment of the variations in radii changes within individual blood vessels (Figure 7). Notably, capillary radius varied along the vessel length across the baseline frames, by 24 ± 28% of the mean resting radius. Consequently, point measurements in vessel calibers - that are widely reported in the literature - do not permit accurate estimation of the microvessel volume changes. Together, the within- and across-vessel radii estimations illustrate the pipeline’s ability to capture spatial variations in the vascular reactivity and relate it to other morphological features (e.g., the distance to the closest labelled neuron).

Vascular graph examples:

A. Baseline variability in vessel diameter estimated by the standard deviation of each vessel’s mean radius across baseline time frames. B. Mean change in the vessel radius induced by optogenetic stimulation. C. Mean change in the vertexwise radius, allowing the visualization of heterogeneity of radius changes within each vessel. D. Distance from each vertex to the closest pyramidal neuron. Each row corresponds to the vascular graph of a different mouse.

Vertex-wise radii along vessel lengths of a sample artery, capillary and venule at baseline vs. post-stimulation:

A. MIP of an artery, vein, and capillary segments before (left) and after (right) optogenetic stimulation with 458nm light at 1.1 mW/mm2. The artery and capillary dilated by 1.33±0.86 μm and 0.42±0.39 μm, respectively (for both p<1e-4, Mann-Whitney U test), whereas there was no significant change in the venular caliber upon photostimulation (p=0.22, Mann-Whitney U test). B. Estimates of the vertex-wise radius obtained along each of the three vessels’ centrelines, before and after stimulation. C. Vertex-wise radii changes in response to optogenetic stimulation. D. The vertex-wise distance from the vascular surface to the closest YFP-expressing neuron.

Vascular reactivity to optogenetic stimulation

Earlier version: The ability of the pipeline to reveal novel spatial relationships between the vascular network reactivity and neuronal activation was demonstrated by examining the relationship between photostimulation-induced microvascular radii responses and 1) the closest YFP-labelled pyramidal neurons within 80μm, and 2) the cortical depth, at the vertex-wise level and across different photostimulations (Figure 8). Vessels were coarsely segregated into small (average radius < 5 μm) and large (average radius > 5 μm) vessels, as we expected them to respond differently due to their differential wall-associated cell composition. Only vessels that significantly responded following optogenetic stimulation were analyzed: a vessel was deemed a responder if its radius changed by more than twice the baseline standard deviation in the vessel’s radius. The morphometric properties of the responders, under different stimulation conditions, are listed in Table 2. The average magnitude of the capillary radius change was 1.04 ± 1.11 μm. The variability in the radius change within the vessel was higher in the dilating vessels, 0.77 ± 0.61 μm, then in the constricting vessels, 0.69 ± 0.49 μm, for 458-nm, 4.3 photostimulation (p<1e-4) (with no statistically significant changes detected for 458-nm, 1.1 photostimulation, and 552-nm, 4.3 photostimulations). Excluding vessels that did not change their radius by over twice the baseline standard deviation removed almost all large vessels from this analysis (Notwithstanding, Supplemental Figure 3 depicts unfiltered large vessel constrictions and dilations.) We ran mixed effects models (at the vessel level) separately on constricting and dilating vessels to investigate the effect of optogenetic stimulation power on the vessel radius changes. Each of the models was run separately on small and large vessels.

Revised version: The ability of the pipeline to reveal novel spatial relationships between the vascular network reactivity and labelled neurons was demonstrated by examining the relationship between photostimulation-induced microvascular radii responses and 1) the closest YFP-labelled pyramidal neurons within 80μm, and 2) the cortical depth, at the vertex-wise level and across different photostimulations (Figure 8). Vessels were coarsely segregated into small (average radius < 5 μm) and large (average radius > 5 μm) vessels, as we expected them to respond differently due to their differential wall-associated cell composition [6], [64], [65], [66], [67], [68], [69], [70], [71]. Only vessels longer than 20 μm (i.e. vessels whose radius was computed by averaging over many cross-sectional planes) that significantly responded following optogenetic stimulation were analyzed: a vessel was deemed a responder if its radius changed by more than twice the baseline standard deviation in the vessel’s radius. The morphometric properties of the responders, under different stimulation conditions, are listed in Table 2. The average magnitude of significant microvascular radius changes across all stimulation conditions was 1.04 ± 1.11 μm (62±47%). The variability in the radius change within the vessel was higher in the dilating vessels, 0.77 ± 0.61 μm (66±72%) than in the constricting vessels, 0.69 ± 0.49 μm (46±18%), for 458-nm, photostimulation (p<1e-4) (with no statistically significant changes detected for either 458-nm, photostimulation or 552-nm, photostimulation).

Optogenetic activation-induced changes in vessel-wise microvascular radii:

Earlier version: Capillary responses included both dilatations, shown in A. and constrictions, shown in B., with changes in the magnitude of the capillary response with increased photostimulation power. * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted. C. Changes to capillary radii are displayed in relation to the closest pyramidal neurons. The proportion of vessels constricting increased with the higher intensity of blue light stimulation, and constrictions tended to occur further away from pyramidal neurons than did dilations. D. Mean cortical depth of responding capillaries showed a tendency for dilators to be closer to the surface and for constrictors to be deeper in the tissue.

Revised version: Capillary responses included both dilatations, shown in A. and constrictions, shown in B., with changes in the magnitude of the capillary response with increased photostimulation power. * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted. C. Probability density function of constrictions and dilations for the 4.3 mW/mm2 photostimulation. D. Changes to capillary radii are displayed in relation to the closest pyramidal neurons. The proportion of vessels constricting increased with the higher intensity of blue light stimulation, and constrictions tended to occur further away from pyramidal neurons than did dilations. E. Mean cortical depth of responding capillaries showed a tendency for dilators to be closer to the surface and for constrictors to be deeper in the tissue.

Added in the revised version: Excluding vessels that did not change their radius by over twice the baseline standard deviation removed almost all large vessels from this analysis. (Notwithstanding, Supplementary Figure 5 depicts unfiltered large vessel constrictions and dilations.) We ran mixed effects models (at the vessel level) separately on constricting and dilating vessels to investigate the effect of optogenetic stimulation power on the vessel radius changes. Each of the models was run separately on small and large vessels.

Removed in the revised version: Microvascular Network Coordination Following Optogenetic Stimulation:

Removed in the revised version: A. Graph representation of a vascular network of 425 vascular segments from a single image stack. Vessel segments are depicted as nodes of the graph; vascular segments that are joined at junctions are connected by edges. Nodes are coloured by the change in the mean vessel-wise radius following photostimulation with 458 nm light at 4.3 mW/mm2. B. Assortativity of photostimulation-induced changes in mean capillary radius increased with increasing photostimulation power. C. Photostimulation-induced changes in the efficiency of the capillary network. The capillary network efficiency changed by a median -0.16 PΩ-1 (IQR: -0.39 to 0.10 PΩ-1) in response to green light; -0.14 PΩ-1 (IQR: -0.55 to 0.27 PΩ-1) in response to lower intensity bluelight; and 0.22 PΩ-1 (IQR = -0.43;1.47 PΩ-1) in response to higher intensity blue light.. There was a significant increase (p = 0.03) in the capillary network efficiency post 458-nm light at 4.3 mW/mm2, when compared to that following the control green illumination. The measurements came from 72 paired acquisitions of 32 image stacks acquired in 17 mice (9M/8F). * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted.

Earlier version: Vessels further away from activated neurons constrict while those closer to the activated neurons dilate

Revised version: Vessels further away from labelled neurons constrict while those closer to the activated neurons dilate

Added in the revised version: We examined the relationship between vascular radius changes and the distance to the closest labelled pyramidal neuron, as microvascular response is thought to result from neuronal activation-elicited generation of vasoactive molecules that diffuse to the neighbouring vessels.

Earlier version: We examined the relationship between vascular radius changes and the distance to the closest labelled pyramidal neuron, as microvascular response is thought to result from neuronal activation-elicited generation of vasoactive molecules that diffuse to the neighbouring vessels. For the control condition (552nm, 4.3 photostimulation), 2.9% of small capillaries dilated while 1.0% of small capillaries constricted; in larger vessels, barely any responded (0.4% dilated). For this control condition, there was no significant difference in the distance from constrictors or dilators to the closest pyramidal neuron. For the 458 nm photostimulation, capillary constrictors were on average farther away than were dilators from the activated pyramidal neuron: dilations occurred 16.8 ± 13.5 μm away from activated neurons while constrictions occurred 22.7 ± 16.3 μm for 1.1 photostimulation (p=1.5e-3) whereas the 4.3 photostimulation had dilations occur 16.1 ± 14.3 μm away and 21.9 ± 14.6 μm for constrictors (p<1e-4). There was no significant shift between the distance from vessels to neurons for the 1.1 and 4.3 stimulations with 458 nm light.. Dilations in capillaries following 458 nm photostimulation were larger than those following the 552 nm control photostimulation: 0.90 ± 0.93 μm dilatations occurred with 1.1 and 0.90 ± 0.78 μm with 4.3 at 458 nm; vs. 0.58 ± 0.92 μm with 4.3 at 552 nm (p<1e-4). For constrictions, 458 nm photostimulations led to -1.39 ± 1.51 μm radius changes with 1.1 and -1.20 ± 1.13 μm radius changes with 4.3 (p=4.4e-3); whereas 552 nm photostimulation induced -0.37 ± 0.30 μm radius changes with 4.3 of power, which was smaller than the 458-nm induced responses (p=0.02).

Revised version: For the control condition (552nm, photostimulation), 2.9% of small capillaries dilated while 1.0% of small capillaries constricted; in larger vessels, barely any responded (0.4% dilated). For this control condition, there was no significant difference in the distance from constrictors or dilators to the closest pyramidal neuron. For the 458 nm photostimulation, capillary constrictors were on average farther away than were dilators from the labelledlabelled pyramidal neuron: dilations occurred 16.8 ± 13.5 μm away from labelledlabelled neurons while constrictions occurred 22.7 ± 16.3 μm for photostimulation (p=1.5e-3) whereas the photostimulation had dilations occur 16.1 ± 14.3 μm away and 21.9 ± 14.6 μm for constrictors (p<1e-4). There was no significant shift between the distance from vessels to neurons for the and stimulations with 458 nm light. Dilations in capillaries following 458 nm photostimulation were larger than those following the 552 nm control photostimulation: 0.90 ± 0.93 μm dilatations occurred with and 0.90 ± 0.78 μm with at 458 nm; vs. 0.58 ± 0.92 μm with at 552 nm (p<1e-4). For constrictions, 458 nm photostimulations led to -1.39 ± 1.51 μm radius changes with and -1.20 ± 1.13 μm radius changes with (p=4.4e-3); whereas 552 nm photostimulation induced -0.37 ± 0.30 μm radius changes with of power, which was smaller than the 458-nm induced responses (p=0.02).

Vascular radius changes at increasing cortical depths

Earlier version: Vascular responses were next segregated by the cortical depth of the vessel (i.e. the average vessel distance from the cortical surface) [64]. Dilators tended to be located closer to the cortical surface across all stimulation conditions. Constricting vessels were located at an average 58 ± 187 μm deeper than dilators for 458-nm stimulation at 1.1 (p=0.02), and 37 ± 179 μm deeper for 458-nm photostimulation at 4.3 (p<1e-4) (with no change in the mean depth of either constricting or dilating vessels with changes in the photostimulation power).

Revised version: Vascular responses were next segregated by the cortical depth of the vessel (i.e. the average vessel distance from the cortical surface) [72]. Dilators tended to be located closer to the cortical surface across all stimulation conditions. Constricting vessels were located at an average 58 ± 187 μm deeper than dilators for 458-nm stimulation at and 37 ± 179 μm deeper for 458-nm photostimulation at (with no change in the mean depth of either constricting or dilating vessels with changes in the photostimulation power).

Vascular Network Coordination Following Optogenetic Stimulation

We examined the coordination of changes in the microvascular network as a whole via assortativity of radius changes and network efficiency changes. The vessel responses were observed to be assortative, i.e., capillaries mirrored the responses in their neighbours. The increases in stimulation power were accompanied by increases in the assortativity of capillary responses: increasing stimulation level resulted in heightened coordination between adjacent capillaries.

The efficiency increased only at the strongest blue photostimulation, i.e., only at this level of stimulation did the resistivity along the average of all of the shortest paths between junctions in the vascular network decrease, resulting in attenuated resistance to flow through the network. The distribution of changes to the efficiency were highly skewed (with a coefficient of skewness of -1.06 for green illumination, 2.92 for lower intensity blue photostimulation, and 4.87 for higher intensity blue photostimulation). The median increase in the efficiency induced by the higher intensity blue photostimulation, of 4% (IQR: -8% to 38%), was significantly higher than the median -6% (IQR=-9% to 4%) efficiency change following the control green illumination

Discussion

Earlier version: Recent studies have demonstrated temporal propagation and coordination in cerebrovascular responses to neuronal activation, whereby arteries dilated after capillary exposure to increased potassium ion concentration [65]; opposing geometric changes have been reported by some studies in capillaries connected by intercapillary tunnelling nanotubes [24]. However, how the effects of these and other mechanisms influence in situ reactivity of the 3D brain capillary network remains unknown. Here, we developed a pipeline for extracting graphs of brain microvascular networks from in situ 2PFM and examine coordination within and across capillaries. Capillary networks and their geometrical changes were imaged via 2PFM during periods of baseline alternated with photostimulation of ChR2 in pyramidal neurons of transgenic mice, and the microvascular mesh was evaluated every 1-1.73 μm. The vascular morphology was then analyzed vertex- and vessel-wise across the entire network. All vessels exhibited significant heterogeneity in caliber changes along their length. Neuronal activation induced both dilations and constrictions of vessels and the incidence of constrictions increased with increasing cortical depth. As the stimulation power increased, the tendency for vessels to change their radius by an amount similar to their neighbours increased. Only the highest photostimulation intensity elicited an increase in the network efficiency. Our findings reveal an intricate level of coordination among brain microvessels and provide a computational analysis platform for interrogating a host of hypotheses on cerebral microvascular reactivity.

Revised version: Recent studies have demonstrated temporal propagation and coordination in cerebrovascular responses to neuronal activation, whereby arteries dilated after capillary exposure to increased potassium ion concentration [73]; opposing geometric changes have been reported by some studies in capillaries connected by intercapillary tunnelling nanotubes [24]. However, how the effects of these and other mechanisms influence in situ reactivity of the 3D brain capillary network remains unknown. Here, we developed a pipeline for extracting graphs of brain microvascular networks from in situ 2PFM and examine coordination within and across capillaries. Capillary networks and their geometrical changes were imaged via 2PFM during periods of baseline alternated with photostimulation of ChR2 in pyramidal neurons of transgenic mice, and the microvascular mesh was evaluated every 1-1.73 μm. The vascular morphology was then analyzed vertex- and vessel-wise across the entire network. All vessels exhibited significant heterogeneity in caliber changes along their length. Neuronal activation induced both dilations and constrictions of vessels and the incidence of constrictions increased with increasing cortical depth. As the stimulation power increased, the tendency for vessels to change their radius by an amount similar to their neighbours increased. Only the highest photostimulation intensity elicited an increase in the network efficiency. Our findings reveal an intricate level of coordination among brain microvessels and provide a computational analysis platform for interrogating a host of hypotheses on cerebral microvascular reactivity.

Vascular Segmentation and Network Extraction

Earlier version: Intensity thresholding-based image processing pipelines have been used to examine vascular networks and quantify vascular morphology, but they have not gained widespread use due to difficulties in adapting them to highly heterogeneous levels of noise across samples [66]– [69]. Deep learning-based methods for analyzing vascular morphology from 3D microscopy images have become prevalent as they provide robust segmentation results over a wide range of signal-to-noise ratios. Recent work has demonstrated steady improvements of segmentation models’ performance with respect to similarity-based metrics (i.e., Dice scores) [70]–[75]; though surface-based metrics may be better predictors of how amenable segmentation outputs will be to subsequent morphological analysis of the microvascular network. Our final UNETR model was selected based on a combination of performance metrics including mean surface distance, Hausdorf 95% distance, and rater evaluation, to maximize the smoothness of the surface of generated segmentation masks and reduce false positive branch points during centerline extraction; thereby leading to higher fidelity rendering of microvascular networks. Graph generation was greatly facilitated by computing the union of the vascular segmentation masks across all time points as it enabled tracing of capillaries that had stalls at the individual time points (since the accompanying loss of the fluorescent label otherwise resulted in graph discontinuities).

Revised version: Intensity thresholding-based image processing pipelines have been used to examine vascular networks and quantify vascular morphology, but they have not gained widespread use due to difficulties in adapting them to highly heterogeneous levels of noise across samples [68], [74], [75], [76]. Deep learning-based methods for analyzing vascular morphology from 3D microscopy images have become prevalent as they provide robust segmentation results over a wide range of signal-to-noise ratios. Recent work has demonstrated steady improvements of segmentation models’ performance with respect to similarity-based metrics (i.e., Dice scores) [77], [78], [79], [80], [81], [82]; though surface-based metrics may be better predictors of how amenable segmentation outputs will be to subsequent morphological analysis of the microvascular network. Our final UNETR model was selected based on a combination of performance metrics including mean surface distance, Hausdorff 95% distance, and rater evaluation, to maximize the smoothness of the surface of generated segmentation masks and reduce false positive branch points during centerline extraction; thereby leading to higher fidelity rendering of microvascular networks. Graph generation was greatly facilitated by computing the union of the vascular segmentation masks across all time points as it enabled tracing of capillaries that had stalls at the individual time points (since the accompanying loss of the fluorescent label otherwise resulted in graph discontinuities).

Added in the revised version: We tested the ability of the pipeline to detect changes to simulated changes to images and the pipeline’s sensitivity to perturbations to extracted vascular morphology. Using resizing, we confirmed that the boundary detection algorithm was able to detect prescribed changes (cf. Figure 5). The average radius estimates for vessels were then shown to vary by 0.64±3.44% upon changes in centerline position, demonstrating that the radius estimates have a low sensitivity to small (2-3 μm) perturbations in centerline placement. Visual examples of vessels repeatedly dilating or constricting are shown in Supplementary Figures 6 and 7.

Earlier version: Network morphological properties at baseline were in line with prior work. The vascular length density measured from fixed tissue ranged from 0.44 to 1.10 m/mm3 [66], [76]–[80]. Our reported vascular density in the forelimb region of the primary somatosensory cortex was 0.40 ± 0.22 m/mm3, with the low end of the range value expected due to fluorescence absorption by hemoglobin in the large pial vessels leading to signal dropout (or shadowing) in the underlying tissue. Our reported average capillary radius of 2.19 ± 1.66 μm was also in line with other studies, where the mean capillary radius ranged from 1.75 to 2.2 μm as measured with confocal microscopy or 2PFM [7], [66].

Revised version: Network morphological properties at baseline were in line with prior work. The vascular length density measured from fixed tissue ranged from 0.44 to 1.10 m/mm3 [74], [83], [84], [85], [86], [87]. Our reported vascular density in the forelimb region of the primary somatosensory cortex was 0.40 ± 0.22 m/mm3, with the low end of the range value expected due to fluorescence absorption by hemoglobin in the large pial vessels leading to signal dropout (or shadowing) in the underlying tissue. Our reported average capillary radius of 2.19 ± 1.66 μm was also in line with other studies, where the mean capillary radius ranged from 1.75 to 2.2 μm as measured with confocal microscopy or 2PFM [7], [74].

In Situ Changes to Vessel Calibers Upon Neuronal Activation

Earlier version: Caliber changes at individual vertices along vessel centerlines exhibited significant heterogeneity. Such heterogeneity is expected due to non-uniformly distributed alpha smooth muscle actin-containing cells along vessel walls, as well as differential activations leading to heterogeneous metabolic demand within the tissue [1], [6], [81]–[84]. Many previous studies assumed vessel caliber changes to be uniform, compromising the accuracy of the estimates.

Revised version: Caliber changes at individual vertices along vessel centerlines exhibited significant heterogeneity. Such heterogeneity is expected due to non-uniformly distributed alpha smooth muscle actin-containing cells along vessel walls, as well as differential activations leading to heterogeneous metabolic demand within the tissue [1], [6], [88], [89], [90], [91]. Many previous studies assumed vessel caliber changes to be uniform, compromising the accuracy of the estimates.

Earlier version: As expected, the control 552 nm stimulation led to minimal changes in vessel calibers. To probe for off-target effects, non-transgenic mice were also tested with the same optical setup and photostimulation, with no changes to vascular diameters observed at any of the photostimulation powers utilized (data not shown). In transgenic mice, we detected an average capillary dilation in significantly responding vessels of 70±83% with low-intensity 458 nm stimulation, and 67±61% with higher intensity 458 nm stimulation. Across photostimulation conditions, the capillary dilations ranged from 2% to 805%. These calibre changes were higher than those previously reported, which varied from 2% to 20% depending on the capillary branch order [6], [7], [23], [85], [86]. Far less data are available on constrictions. In the current work, constrictions averaged 47±20% for lower-intensity blue light stimulation and 47±17% for higher-intensity blue light stimulation, with a constriction range of 5% to 97% of the baseline radius. These are higher than the previously reported constrictions of 20% [6], [23], likely due to our identifying as responding vessels only those whose caliber changed by at least twice their baseline caliber variation.

Revised version: As expected, the control 552 nm stimulation led to minimal changes in vessel calibers. To probe for off-target effects, non-transgenic mice were also tested with the same optical setup and photostimulation, with no changes to vascular diameters observed at any of the photostimulation powers utilized (cf. Supplementary Figure 8). In transgenic mice, we detected an average capillary dilation in significantly responding vessels of 70±83% with low-intensity 458 nm stimulation, and 67±61% with higher intensity 458 nm stimulation. Across photostimulation conditions, the capillary dilations ranged from 2% to 805%. These calibre changes were higher than those previously reported, which varied from 2% to 20% depending on the capillary branch order [6], [7], [23], [92], [93]. Far less data are available on constrictions. In the current work, constrictions averaged 47±20% for lower-intensity blue light stimulation and 47±17% for higher-intensity blue light stimulation, with a constriction range of 5% to 97% of the baseline radius. These are higher than the previously reported constrictions of 20% [6], [23], likely due to our identifying as responding vessels only those whose caliber changed by at least twice their baseline caliber variation. It is also worth noting that vessels’ response directions were consistent on repeated trials. Of the vessels whose radius change exceeded twice the baseline variability across time, 31.7% dilated on some trials while constricting on others; 41.1% dilated on each trial; and 27.2% constricted on each trial. (Note that some trials use 1.1 vs. 4.3 mW/mm2 and some have opposite scanning directions).

Earlier version: 458-nm photostimulation caused capillaries to dilate when pyramidal neurons were close, and constrict when they were further away. However, the stronger blue light stimulation led to an increased rate of constrictions, double that of the low-powered blue light stimulation. For larger vessels, both 458 nm stimulation powers led to a similar dilation level that diminished with increasing distance from activated pyramidal neurons. This tendency for vessels close to neurons to dilate and further away ones to constrict would be expected in flow redirection into regions of high level of neuronal activity. Stimulation power dependence in blood flow changes has previously been reported in optogenetic mouse models with diffuse stimulation via LED probes, and following transcranial alternating current stimulation [87], [88]. However, neither of the previously employed methods was able to discern the spatial relationship between the vascular caliber changes, or relate these changes to the distribution of the stimulated neurons.

Revised version: 458-nm photostimulation resulted in a mix of constrictions and dilations with 44.1% of significantly responding vessels within 10 μm of a labelled pyramidal neuron constricting and 55.1% dilating, while 53.3% of vessels further than 30 μm constricted and 46.7% dilated. The cutoff distances from the closest labelled neuron were based on estimates of cerebral metabolic rate of oxygen consumption that showed a steep gradient in oxygen consumption with distance from arteries, CMRO2 being halved by 30 μm away [94]. The stronger blue light stimulation led to an increased rate of constrictions, double that of the low-powered blue light stimulation. For larger vessels, both 458 nm stimulation powers led to a similar dilation level that diminished with increasing distance from labelled pyramidal neurons. This tendency for vessels close to neurons to dilate and further away ones to constrict would be expected in flow redirection into regions of high level of neuronal activity. Stimulation power dependence in blood flow changes has previously been reported in optogenetic mouse models with diffuse stimulation via LED probes, and following transcranial alternating current stimulation [95], [96]. However, neither of the previously employed methods was able to discern the spatial relationship between the vascular caliber changes, or relate these changes to the distribution of the stimulated neurons.

Earlier version: Across all stimulations, constricting vessels tended to be located deeper than dilating vessels. As the blue light stimulation power increased, the mean depth of both constricting and dilating vessels increased, likely resulting from higher intensity light reaching pyramidal neurons deeper in the tissue [89], [90]. Our results underscore that the haemodynamic response following targeted neuronal activation is not uniformly distributed across the microvascular network: accurate neurovascular coupling assessment thus requires network-based analysis.

Revised version: As the blue light stimulation power increased, the mean depth of both constricting and dilating vessels increased, likely resulting from higher intensity light reaching pyramidal neurons deeper in the tissue [97], [98]. The blue light would be expected to excite a lower number of neurons farther from the cortical surface at lower powers. Our results underscore that the haemodynamic response following targeted neuronal activation is not uniformly distributed across the microvascular network: accurate neurovascular coupling assessment thus requires network-based analysis.

Vascular Network Reactivity

Earlier version: To study the microvascular network response as a whole, we examined the assortativity between capillary radius changes and network efficiency changes following optogenetic stimulation. These two graph theory metrics were selected as they both leverage the knowledge of the vascular network structure. Assortativity sheds light on how the vascular network coordinates its responses, while efficiency provides insight into the extent to which those changes facilitate flow through the network. The assortativity revealed that as the stimulation power increased, the tendency of vessels to match their changes to those of their neighbours increased. Previously characterized assortative mechanisms include endothelial cell cation conduction via Kir2.1 channels to synchronize vascular responses [65], and spatial adjacency of pericytes on in vitro retinal preparation leading to assortative changes in neighbouring capillaries [81]. Disassortative (causing opposite changes) mechanisms of capillary coordination have also previously been observed in situ and may result from intercapillary nanotubules’ signalling causing connected pericytes to undergo opposing changes [24]. While not ruling out the presence of disassortative control mechanisms, our results suggest that assortative mechanisms dominate capillary responses to neuronal activation in the somatosensory cortex.

Revised version: To study the microvascular network response as a whole, we examined the assortativity between capillary radius changes and network efficiency changes following optogenetic stimulation. These two graph theory metrics were selected as they both leverage the knowledge of the vascular network structure. Assortativity sheds light on how the vascular network coordinates its responses, while efficiency provides insight into the extent to which those changes facilitate flow through the network. The assortativity revealed that as the stimulation power increased, the tendency of vessels to match their changes to those of their neighbours increased. Previously characterized assortative mechanisms include endothelial cell cation conduction via Kir2.1 channels to synchronize vascular responses [73], and spatial adjacency of pericytes on in vitro retinal preparation leading to assortative changes in neighbouring capillaries [88]. Disassortative (causing opposite changes) mechanisms of capillary coordination have also previously been observed in situ and may result from intercapillary nanotubules’ signalling causing connected pericytes to undergo opposing changes [24]. While not ruling out the presence of disassortative control mechanisms, our results suggest that assortative mechanisms dominate capillary responses to neuronal activation in the somatosensory cortex.

Earlier version: The network efficiency here can be thought of as paralleling mean transit time, i.e., the time it takes blood to traverse the capillary network from the arteries to the veins. In situ studies of mean transit time have revealed a high heterogeneity of plasma traversal of the capillary bed during stimulation, with stimulation reducing plasma transit times by 11% to 20% from its resting levels [86], [91], and simulations suggesting that capillary network geometry and locations of caliber changes exert a substantial influence on these responses [92]. The efficiency of the vascular network here increased significantly only with the strongest 458 nm stimulation. Small dilatations may thus not increase flow in the cortex. The differences in efficiency are likely due to the patterns of localized dilations and constrictions within the vascular network. Efficiency calculations are sensitive to bottlenecks when traversing meshes and certain locations constricting or dilating can have profound impacts on the shortest paths between nodes and the path’s resistivity. The highest powered 458 nm stimulation increasing efficiency may have resulted from increased assortativity causing dilation in key locations within the microvascular network, leading to a significant reduction in shortest path resistivity.

Revised version: The network efficiency here can be thought of as paralleling mean transit time, i.e., the time it takes blood to traverse the capillary network from the arteries to the veins. In situ studies of mean transit time have revealed a high heterogeneity of plasma traversal of the capillary bed during stimulation, with stimulation reducing plasma transit times by 11% to 20% from its resting levels [93], [99], and simulations suggesting that capillary network geometry and locations of caliber changes exert a substantial influence on these responses [100]. The efficiency of the vascular network here increased significantly only with the strongest 458 nm stimulation. Small dilatations may thus not increase flow in the cortex. The differences in efficiency are likely due to the patterns of localized dilations and constrictions within the vascular network. Efficiency calculations are sensitive to bottlenecks when traversing meshes and certain locations constricting or dilating can have profound impacts on the shortest paths between nodes and the path’s resistivity. The highest powered 458 nm stimulation increasing efficiency may have resulted from increased assortativity causing dilation in key locations within the microvascular network, leading to a significant reduction in shortest path resistivity.

Pipeline Limitations and Adaptability

Earlier version: The segmentation model was trained only on vascular and neuronal labels, limiting its generalizability to segmenting alternative cells in the current state. However, it can easily be fine-tuned or retrained to label other brain cells (e.g., pericytes, astrocytes, or endothelial cells). In addition, the segmentation model performed poorly when significant bleeding occurred in the cranial window, compromising the vascular contrast. Our imaging protocol was challenged by the need to resolve individual vessel responses and capture the entire network within the span of the microvascular response to stimulation: we could thus not comment on the temporal evolution of the microvascular response. The temporal evolution of the vascular response in individual vessels, however, has been reported on already using line scanning acquisitions to measure red blood cell velocity and flux [6], [9], [23], [30], [32], [86]. Moreover, the present analysis pipeline can readily be applied to 2PFM data obtained with finer temporal (e.g., via a Piezo objective positioner) or spatial resolution, and/or larger fields of view. Finally, microvascular networks in different brain areas may show distinct spatiotemporal profiles of response to neuronal activation. Future work is required to test the generalizability of present findings across different brain regions.

Revised version: The segmentation model was trained only on vascular and neuronal labels, limiting its generalizability to segmenting alternative cells in the current state. However, it can easily be fine-tuned or retrained to label other brain cells (e.g., pericytes, astrocytes, or endothelial cells). Our vascular segmentation model generalized well to C57BL/6 mouse and Fischer rat data, as well as to Thy1-ChR2 light-sheet fluorescence microscopy images gathered on an UltraMicroscope Blaze lightsheet fluorescence microscope (Miltenyi Biotech) (cf. Supplementary Figure 9 and Supplementary Table 3). However, the segmentation model performed poorly when significant bleeding occurred in the cranial window, compromising the vascular contrast. Our imaging protocol was challenged by the need to resolve individual vessel responses and capture the entire network within the span of the microvascular response to stimulation: thus, we could not examine the kinetics of the microvascular response. The temporal evolution of the response in individual vessels, however, has been reported on using line scanning acquisitions to measure red blood cell velocity and flux [6], [9], [23], [30], [32], [93]. Moreover, the present analysis pipeline can readily be applied to 2PFM data obtained with finer temporal (e.g., via a Piezo objective positioner) or spatial resolution, and/or different size fields of view. Additionally, alternative definitions of responding vessels may be useful depending on the end goal of a study (e.g. selecting a threshold for the radius change based on a percentage change from the baseline level: cf. Supplementary Figure 10 for capillary changes above 10% of the baseline radius). Finally, microvascular networks in different brain areas may show distinct spatiotemporal profiles of response to neuronal activation. Future work is required to test the generalizability of present findings across different brain regions.

Conclusion

We developed a novel deep learning-based computational pipeline for analysis of a time series of 3D 2PFM images and investigation of spatial patterns in microvascular network reactivity to neuronal activation. The microvascular network was represented as a graph, allowing for the evaluation of network geometry changes over time. We tracked the size of blood vessels throughout the network and related vessel radius changes to the distance from the stimulated neurons and the cortical depth. Neuronal activation induced both dilatations and constrictions of capillaries, and the magnitude of these responses increased with increased photostimulation levels while showing significant heterogeneity within and between vessels. In the analysis presented, vertex-wise measurements were aggregated for vessel-wise analysis resulting in highly robust estimates of vessels’ calibers and allowing ready comparisons to literature. Notwithstanding, the pipeline also affords vertex-wise analysis and thus registration of microvascular reactivity with other local morphological features, at an unprecedented spatial scale. With increasing distance of the vessel from the most proximal activated neuron, dilatation magnitude decreased and the incidence of constrictions increased. At the highest stimulation level investigated, the incidence of vessel constrictions also increased with cortical depth. With increasing activation levels, capillaries displayed diameter changes that were similar to their immediate neighbours, while vascular network efficiency increased only under the strongest stimulation. Our computational analysis pipeline permits probing microvascular network reactivity and sheds light on the heterogeneity and coordination of vessel caliber changes across the microvascular network.The pipeline will be made available to the research community to propel future studies of neurovascular coupling and network reactivity.